Brain-derived estrogens facilitate male-typical behaviors by potentiating androgen receptor signaling in medaka

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

This research group has consistently performed cutting-edge research aiming to understand the role of hormones in the control of social behaviors, specifically by utilizing the genetically-tractable teleost fish, medaka, and the current work is no exception. The overall claim they make, that estrogens modulate social behaviors in males and females is supported, with important caveats. For one, there is no evidence these estrogens are generated by "neurons" as would be assumed by their main claim that it is NEUROestrogens that drive this effect. While indeed the aromatase they have investigated is expressed solely in the brain, in most teleosts, brain aromatase is only present in glial cells (astrocytes, radial glia). The authors should change this description so as not to mislead the reader. Below I detail more specific strengths and weaknesses of this manuscript.

We thank the reviewer for this positive evaluation of our work and for the helpful comments and suggestions. Regarding the concern that the term “neuroestrogens” may be misleading, we addressed this in the previous revision by consistently replacing it throughout the manuscript with “brain-derived estrogens” or “brain estrogens.”

In addition, the following sentence was added to the Introduction (line 61): “In teleost brains, including those of medaka, aromatase is exclusively localized in radial glial cells, in contrast to its neuronal localization in rodent brains (Forlano et al., 2001; Diotel et al., 2010; Takeuchi and Okubo, 2013).”

Strenghth:

Excellent use of the medaka model to disentangle the control of social behavior by sex steroid hormones

The findings are strong for the most part because deficits in the mutants are restored by the molecule (estrogens) that was no longer present due to the mutation

Presentation of the approach and findings are clear, allowing the reader to make their own inferences and compare them with the authors'

Includes multiple follow-up experiments, which leads to tests of internal replication and an impactful mechanistic proposal

Findings are provocative not just for teleost researchers, but for other species since, as the authors point out, the data suggest mechanisms of estrogenic control of social behaviors may be evolutionary ancient

We thank the reviewer again for their positive evaluation of our work.

Weakness:

As stated in the summary, the authors are attributing the estrogen source to neurons and there isn't evidence this is the case. The impact of the findings doesn't rest on this either

As mentioned above, we addressed this in the previous revision by replacing “neuroestrogens” with “brain-derived estrogens” or “brain estrogens” throughout the manuscript. In addition, the following sentence was added to the Introduction (line 61): “In teleost brains, including those of medaka, aromatase is exclusively localized in radial glial cells, in contrast to its neuronal localization in rodent brains (Forlano et al., 2001; Diotel et al., 2010; Takeuchi and Okubo, 2013).”

The d4 versus d8 esr2a mutants showed different results for aggression. The meaning and implications of this finding are not discussed, leaving the reader wondering

This comment is the same as one raised in the first review (Reviewer #1’s comment 2 on weaknesses), which we already addressed in our initial revision. For the reviewer’s convenience, we provide the response below:

Line 300: As the reviewer correctly noted, circles were significantly reduced in mutant males of the Δ8 line, whereas no significant reduction was observed in those of the Δ4 line. However, a tendency toward reduction was evident in the Δ4 line (P = 0.1512), and both lines showed significant differences in fin displays. Based on these findings, we believe our conclusion that esr2a−/− males exhibit reduced aggression remains valid. To clarify this point and address potential reader concerns, we have revised the text as follows: “esr2a−/− males exhibited significantly fewer fin displays (P = 0.0461 and 0.0293 for Δ8 and Δ4 lines, respectively) and circles (P = 0.0446 and 0.1512 for Δ8 and Δ4 lines, respectively) than their wild-type siblings (Fig. 5L; Fig. S8E), suggesting less aggression” was edited to read “esr2a−/− males from both the Δ8 and Δ4 lines exhibited significantly fewer fin displays than their wild-type siblings (P = 0.0461 and 0.0293, respectively). Circles followed a similar pattern, with a significant reduction in the Δ8 line (P = 0.0446) and a comparable but non-significant decrease in the Δ4 line (P =0.1512) (Figure 5L, Figure 5—figure supplement 3E), showing less aggression.”

Lack of attribution of previous published work from other research groups that would provide the proper context of the present study

This comment is also the same as one raised in the first review (Reviewer #1’s comment 3 on weaknesses). In our previous revision, in response to this comment, we cited the relevant references (Hallgren et al., 2006; O’Connell and Hofmann, 2012; Huffman et al., 2013; Jalabert et al., 2015; Yong et al., 2017; Alward et al., 2020; Ogino et al., 2023) in the appropriate sections. We also added the following new references and revised the Introduction and Discussion accordingly:

(2) Alward BA, Laud VA, Skalnik CJ, York RA, Juntti SA, Fernald RD. 2020. Modular genetic control of social status in a cichlid fish. Proceedings of the National Academy of Sciences of the United States of America 117:28167–28174. DOI: https://doi.org/10.1073/pnas.2008925117

(39) O’Connell LA, Hofmann HA. 2012. Social status predicts how sex steroid receptors regulate complex behavior across levels of biological organization. Endocrinology 153:1341–1351. DOI:https://doi.org/10.1210/en.2011-1663

(54) Yong L, Thet Z, Zhu Y. 2017. Genetic editing of the androgen receptor contributes to impaired male courtship behavior in zebrafish. Journal of Experimental Biology 220:3017–3021.DOI:https://doi.org/10.1242/jeb.161596

There are a surprising number of citations not included; some of the ones not included argue against the authors' claims that their findings were "contrary to expectation"

In our previous revision, we cited the relevant references (Hallgren et al., 2006; O’Connell and Hofmann, 2012; Huffman et al., 2013; Jalabert et al., 2015) in the Introduction. We also revised the text to remove phrases such as “contrary to expectation” and “unexpected.”

The experimental design for studying aggression in males has flaws. A standard test like a residentintruder test should be used.

Following this comment, we have attempted additional aggression assays using the resident-intruder paradigm. However, these experiments did not produce consistent or interpretable results. As noted in our previous revision, medaka naturally form shoals and exhibit weak territoriality, and even slight differences in dominance between a resident and an intruder can markedly increase variability, reducing data reliability. Therefore, we believe that the approach used in the present study provides a more suitable assessment of aggression in medaka, regardless of territorial tendencies. We will continue to explore potential refinements in future studies and respectfully ask the reviewer to evaluate the present work based on the assay used here.

While they investigate males and females, there are fewer experiments and explanations for the female results, making it feel like a small addition or an aside

While we did not adopt this comment in our previous revision, we have carefully reconsidered the reviewers’ feedback and have now decided to remove the female data. This change allows us to present a more focused and cohesive story centered on males. The specific revisions are outlined below:

Abstract

Line 25: The text “, thereby revealing a previously unappreciated mode of action of brain-derived estrogens. We additionally show that female fish lacking Cyp19a1b are less receptive to male courtship and conversely court other females, highlighting the significance of brain-derived estrogens in establishing sex-typical behaviors in both sexes.” has been revised to “. Taken together, these findings reveal a previously unappreciated mode of action of brain-derived estrogens in shaping male-typical behaviors.”

Results

Line 88: The text “Loss of cyp19a1b function in these fish was verified by measuring brain and peripheral levels of sex steroids. As expected, brain estradiol-17β (E2) in both male and female homozygous mutants (cyp19a1b−/−) was significantly reduced to 16% and 50%, respectively, of the levels in their wild-type (cyp19a1b+/+) siblings (P = 0.0037, males; P = 0.0092, females) (Fig. 1, A and B). In males, brain E2 in heterozygotes (cyp19a1b−/−) was also reduced to 45% of the level in wild-type siblings (P = 0.0284) (Fig. 1A), indicating a dosage effect of cyp19a1b mutation. In contrast, peripheral E2 levels were unaltered in both cyp19a1b−/− males and females (Fig. S1, C and D), consistent with the expected functioning of Cyp19a1b primarily in the brain. Strikingly, brain levels of testosterone, as opposed to E2, increased 2.2-fold in cyp19a1b−/− males relative to wild-type siblings (P = 0.0006) (Fig. 1A). Similarly, brain 11KT levels in cyp19a1b−/− males and females increased 6.2- and 1.9-fold, respectively, versus wild-type siblings (P = 0.0007, males; P = 0.0316, females) (Fig. 1, A and B). These results show that cyp19a1b-deficient fish have reduced estrogen levels coupled with increased androgen levels in the brain, confirming the loss of cyp19a1b function. They also suggest that the majority of estrogens in the male brain and half of those in the female brain are synthesized locally in the brain. In addition, peripheral 11KT levels in cyp19a1b−/− males and females increased 3.7- and 1.8-fold, respectively (P = 0.0789, males; P = 0.0118, females) (Fig. S1, C and D), indicating peripheral influence in addition to central effects.” has been revised to “Loss of cyp19a1b function in these fish was verified by measuring brain and peripheral levels of sex steroids in males. As expected, brain estradiol-17β (E2) in homozygous mutants (cyp19a1b−/−) was significantly reduced to 16% of the levels in wild-type (cyp19a1b+/+) siblings (P = 0.0037) (Figure 1A). Brain E2 in heterozygotes (cyp19a1b+/−) was also reduced to 45% of wild-type levels (P = 0.0284) (Figure 1A), indicating a dosage effect of the cyp19a1b mutation. In contrast, peripheral E2 levels were unaltered in cyp19a1b−/− males (Figure 1B), consistent with the expected functioning of Cyp19a1b primarily in the brain. Strikingly, brain testosterone levels, as opposed to E2, increased 2.2-fold in cyp19a1b−/− males relative to wild-type siblings (P = 0.0006) (Figure 1A). Similarly, brain 11KT levels increased 6.2-fold (P = 0.0007) (Figure 1A). These results indicate that cyp19a1b-deficient males have reduced estrogen coupled with elevated androgen levels in the brain, confirming the loss of cyp19a1b function. They also suggest that the majority of estrogens in the male brain are synthesized locally in the brain. Peripheral 11KT levels also increased 3.7-fold in cyp19a1b−/− males (P = 0.0789) (Figure 1B), indicating peripheral influence in addition to central effects.”

Line 211: “expression of vt in the pNVT of cyp19a1b−/− males was significantly reduced to 18% as compared with cyp19a1b+/+ males (P = 0.0040), a level comparable to that observed in females” has been revised to “expression of vt in the pNVT of cyp19a1b−/− males was significantly reduced to 18% as compared with cyp19a1b+/+ males (P = 0.0040).”

The subsection entitled “cyp19a1b-deficient females are less receptive to males and instead court other females,” which followed line 311, has been removed.

Discussion

The two paragraphs between lines 373 and 374, which addressed the female data, have been removed.

Materials and methods

Line 433: “males and females” has been changed to “males”.

Line 457: “focal fish” has been changed to “focal male”.

Line 458: “stimulus fish” has been changed to “stimulus female”.

Line 458: “Fig. 6, E and F, ” has been deleted.

Line 460: “; wild-type males in Fig. 6, A to C” has been deleted.

Line 466: The text “The period of interaction/recording was extended to 2 hours in tests of courtship displays received from the stimulus esr2b-deficient female and in tests of mating behavior between females, because they take longer to initiate courtship (12). In tests using an esr2b-deficient female as the stimulus fish, where the latency to spawn could not be calculated because these fish were unreceptive to males and did not spawn, the sexual motivation of the focal fish was instead assessed by counting the number of courtship displays and wrapping attempts in 30 min. The number of these mating acts was also counted in tests to evaluate the receptivity of females. In tests of mating behavior between two females, the stimulus female was marked with a small notch in the caudal fin to distinguish it from the focal female.” has been revised to “In tests using an esr2b-deficient female as the stimulus fish, the latency to spawn could not be calculated because the female was unreceptive to males and did not spawn. Therefore, the sexual motivation of the focal male was assessed by counting the number of courtship displays and wrapping attempts in 30 min. To evaluate courtship displays performed by stimulus esr2bdeficient females toward focal males, the recording period was extended to 2 hours, as these females take longer to initiate courtship (Nishiike et al., 2021). In all video analyses, the researcher was blind to the fish genotype and treatment.”

Line 499: “brains dissected from males and females of the cyp19a1b-deficient line (analysis of ara, arb, vt, gal, npba, and esr2b) and males of the esr1-, esr2a-, and esr2b-deficient lines” has been revised to “male brains from the cyp19a1b-deficient line (analysis of ara, arb, vt, and gal) and from the esr1-, esr2a-, and esr2b-deficient lines.”

Line 504: “After color development for 15 min (gal), 40 min (npba), 2 hours (vt), or overnight (ara, arb, and esr2b)” has been revised to “After color development for 15 min (gal), 2 hours (vt), or overnight (ara and arb).”

Line 516: “Thermo Fisher Scientific, Waltham, MA” has been changed to “Thermo Fisher Scientific” to avoid redundancy.

Line 565: The subsection entitled “Measurement of spatial distances between fish” has been removed.

Line 585: “6/10 cyp19a1b+/+, 3/10 cyp19a1b+/−, and 6/10 cyp19a1b−/− females were excluded in Fig. 6B;” has been deleted.

References

The following references have been removed:

Capel B. 2017. Vertebrate sex determination: evolutionary plasticity of a fundamental switch. Nature Reviews Genetics 18:675–689. DOI: https://doi.org/10.1038/nrg.2017.60

Hiraki T, Nakasone K, Hosono K, Kawabata Y, Nagahama Y, Okubo K. 2014. Neuropeptide B is femalespecifically expressed in the telencephalic and preoptic nuclei of the medaka brain. Endocrinology 155:1021–1032. DOI: https://doi.org/10.1210/en.2013-1806

Juntti SA, Hilliard AT, Kent KR, Kumar A, Nguyen A, Jimenez MA, Loveland JL, Mourrain P, Fernald RD. 2016. A neural basis for control of cichlid female reproductive behavior by prostaglandin F2α. Current Biology 26:943–949. DOI: https://doi.org/10.1016/j.cub.2016.01.067

Kimchi T, Xu J, Dulac C. 2007. A functional circuit underlying male sexual behaviour in the female mouse brain. Nature 448:1009–1014. DOI: https://doi.org/10.1038/nature06089

Kobayashi M, Stacey N. 1993. Prostaglandin-induced female spawning behavior in goldfish (Carassius auratus) appears independent of ovarian influence. Hormones and Behavior 27:38–55.

DOI:https://doi.org/10.1006/hbeh.1993.1004

Liu H, Todd EV, Lokman PM, Lamm MS, Godwin JR, Gemmell NJ. 2017. Sexual plasticity: a fishy tale. Molecular Reproduction and Development 84:171–194. DOI: https://doi.org/10.1002/mrd.22691

Munakata A, Kobayashi M. 2010. Endocrine control of sexual behavior in teleost fish. General and Comparative Endocrinology 165:456–468. DOI: https://doi.org/10.1016/j.ygcen.2009.04.011

Nugent BM, Wright CL, Shetty AC, Hodes GE, Lenz KM, Mahurkar A, Russo SJ, Devine SE, McCarthy MM. 2015. Brain feminization requires active repression of masculinization via DNA methylation. Nature Neuroscience 18:690–697. DOI: https://doi.org/10.1038/nn.3988

Shaw K, Therrien M, Lu C, Liu X, Trudeau VL. 2023. Mutation of brain aromatase disrupts spawning behavior and reproductive health in female zebrafish. Frontiers in Endocrinology 14:1225199.

DOI:https://doi.org/10.3389/fendo.2023.1225199

Stacey NE. 1976. Effects of indomethacin and prostaglandins on the spawning behaviour of female goldfish. Prostaglandins 12:113–126. DOI: https://doi.org/10.1016/s0090-6980(76)80010-x

Figure 1

Panel B, which originally showed steroid levels in female brains, has been replaced with steroid levels in the periphery of males, originally presented in Figure S1, panel C. Accordingly, the legend “(A and B) Levels of E2, testosterone, and 11KT in the brain of adult cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (A) and females (B) (n = 3 per genotype and sex).” has been revised to “(A, B) Levels of E2, testosterone, and 11KT in the brain (A) and periphery (B) of adult cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 3 per genotype).”

Figure 3

The female data have been deleted from Figure 3. The revised Figure 3 is presented.

The corresponding legend text has been revised as follows:

Line 862: “males and females (n = 4 and 5 per genotype for males and females, respectively)” has been changed to “males (n = 4 per genotype)”.

Line 864: “males and females (n = 4 except for cyp19a1b+/+ males, where n = 3)” has been changed to “males (n = 3 and 4, respectively)”.

Figure 6

Figure 6 and its legend have been removed.

Figure 1—figure supplement 1

Panel C, showing male data, has been moved to Figure 1B, as described above, while panel D, showing female data, has been deleted. The corresponding legend “(C and D) Levels of E2, testosterone, and 11KT in the periphery of adult cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (C) and females (D) (n = 3 per genotype and sex). Statistical differences were assessed by Bonferroni’s post hoc test (C and D). Error bars represent SEM. *P < 0.05.” has also been removed.

Line 804: Following this change, the figure title has been updated from “Generation of cyp19a1bdeficient medaka and evaluation of peripheral sex steroid levels” to “Generation of cyp19a1b-deficient medaka.”

The statistics comparing "experimental to experimental" and "control to experimental" isn't appropriate

This comment is the same as one raised in the first review (Reviewer #1’s comment 7 on weaknesses), which we already addressed in our initial revision. For the reviewer’s convenience, we provide the response below:

The reviewer raised concerns about the statistical analysis used for Figures 4C and 4E, suggesting that Bonferroni’s test should be used instead of Dunnett’s test. However, Dunnett’s test is commonly used to compare treatment groups to a reference group that receives no treatment, as in our study. Since we do not compare the treated groups with each other, we believe Dunnett’s test is the most appropriate choice.

Line 576: The reviewer’s concern may have arisen from the phrase “comparisons between control and experimental groups” in the Materials and methods. We have revised it to “comparisons between untreated and E2-treated groups in Figure 4C and D” for clarity.

Reviewer #3 (Public Review):

Summary:

Taking advantage of the existence in fish of two genes coding for estrogen synthase, the enzyme aromatase, one mostly expressed in the brain (Cyp19a1b) and the other mostly found in the gonads (Cyp19a1a), this study investigates the role of brain-derived estrogens in the control of sexual and aggressive behavior in medaka. The constitutive deletion of Cyp19a1b markedly reduced brain estrogen content in males and to a lesser extent in females. These effects are accompanied by reduced sexual and aggressive behavior in males and reduced preference for males in females. These effects are reversed by adult treatment with supporting a role for estrogens. The deletion of Cyp19a1b is associated with a reduced expression of the genes coding for the two androgen receptors, ara and arb, in brain regions involved in the regulation of social behavior. The analysis of the gene expression and behavior of mutants of estrogen receptors indicates that these effects are likely mediated by the activation of the esr1 and esr2a isoforms. These results provide valuable insight into the role of estrogens in social behavior in the most abundant vertebrate taxon, however the conclusion of brain-derived estrogens awaits definitive confirmation.

We thank this reviewer for their positive evaluation of our work and comments that have improved the manuscript.

Strength:

Evaluation of the role of brain "specific" Cyp19a1 in male teleost fish, which as a taxon are more abundant and yet proportionally less studied that the most common birds and rodents. Therefore, evaluating the generalizability of results from higher vertebrates is important. This approach also offers great potential to study the role of brain estrogen production in females, an understudied question in all taxa.

Results obtained from multiple mutant lines converge to show that estrogen signaling, likely synthesized in the brain drives aspects of male sexual behavior.

The comparative discussion of the age-dependent abundance of brain aromatase in fish vs mammals and its role in organization vs activation is important beyond the study of the targeted species. - The authors have made important corrections to tone down some of the conclusions which are more in line with the results.

We thank the reviewer again for their positive evaluation of our work and the revisions we have made.

weaknesses:

No evaluation of the mRNA and protein products of Cyp19a1b and ESR2a are presented, such that there is no proper demonstration that the mutation indeed leads to aromatase reduction. The conclusion that these effects dependent on brain derived estrogens is therefore only supported by measures of E2 with an EIA kit that is not validated. No discussion of these shortcomings is provided in the discussion thus further weakening the conclusion manuscript.

In response to this and other comments, we have now provided direct validation that the cyp19a1b mutation in our medaka leads to loss of function. Real-time PCR analysis showed that cyp19a1b transcript levels in the brain were reduced by approximately half in cyp19a1b+/− males and were nearly absent in cyp19a1b−/− males, consistent with nonsense-mediated mRNA decay

In addition, AlphaFold 3-based structural modeling indicated that the mutant Cyp19a1b protein lacks essential motifs, including the aromatic region and heme-binding loop, and exhibits severe conformational distortion (see figure; key structural features are annotated as follows: membrane helix (blue), aromatic region (red), and heme-binding loop (orange)).

Results:

Line 101: The following text has been added: “Loss of cyp19a1b function was further confirmed by measuring cyp19a1b transcript levels in the brain and by predicting the three-dimensional structure of the mutant protein. Real-time PCR revealed that transcript levels were reduced by half in cyp19a1b+/− males and were nearly undetectable in cyp19a1b−/− males, presumably as a result of nonsense-mediated mRNA decay (Lindeboom et al., 2019) (Figure 1C). The wild-type protein, modeled by AlphaFold 3, exhibited a typical cytochrome P450 fold, including the membrane helix, aromatic region, and hemebinding loop, all arranged in the expected configuration (Figure 1—figure supplement 1C). The mutant protein, in contrast, was severely truncated, retaining only the membrane helix (Figure 1—figure supplement 1C). The absence of essential domains strongly indicates that the allele encodes a nonfunctional Cyp19a1b protein. Together, transcript and structural analyses consistently demonstrate that the mutation generated in this study causes a complete loss of cyp19a1b function.”

Materials and methods

Line 438: A subsection entitled “Real-time PCR” has been added. The text of this subsection is as follows: “Total RNA was isolated from the brains of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males using the RNeasy Plus Universal Mini Kit (Qiagen, Hilden, Germany). cDNA was synthesized with the SuperScript VILO cDNA Synthesis Kit (Thermo Fisher Scientific, Waltham, MA). Real-time PCR was performed on the LightCycler 480 System II using the LightCycler 480 SYBR Green I Master (Roche Diagnostics). Melting curve analysis was conducted to verify that a single amplicon was obtained in each sample. The β-actin gene (actb; GenBank accession number NM_001104808) was used to normalize the levels of target transcripts. The primers used for real-time PCR are shown in Supplementary file 2.”

Line 448: A subsection entitled “Protein structure prediction” has been added. The text of this subsection is as follows: “Structural predictions of Cyp19a1b proteins were conducted using AlphaFold 3 (Abramson et al., 2024). Amino acid sequences corresponding to the wild-type allele and the mutant allele generated in this study were submitted to the AlphaFold 3 prediction server. The resulting models were visualized with PyMOL (Schrödinger, New York, NY), and key structural features, including the membrane helix, aromatic region, and heme-binding loop, were annotated.”

References

The following two references have been added:

Abramson J, Adler J, Dunger J, Evans R, Green T, Pritzel A, Ronneberger O, Willmore L, Ballard AJ, Bambrick J, Bodenstein SW, Evans DA, Hung CC, O'Neill M, Reiman D, Tunyasuvunakool K, Wu Z, Žemgulytė A, Arvaniti E, Beattie C, Bertolli O, Bridgland A, Cherepanov A, Congreve M, CowenRivers AI, Cowie A, Figurnov M, Fuchs FB, Gladman H, Jain R, Khan YA, Low CMR, Perlin K, Potapenko A, Savy P, Singh S, Stecula A, Thillaisundaram A, Tong C, Yakneen S, Zhong ED, Zielinski M, Žídek A, Bapst V, Kohli P, Jaderberg M, Hassabis D, Jumper JM. 2024. Accurate structure prediction of biomolecular interactions with AlphaFold 3. Nature 630:493–500. DOI: https://doi.org/10.1038/s41586-024-07487-w

Lindeboom RGH, Vermeulen M, Lehner B, Supek F. 2019. The impact of nonsense-mediated mRNA decay on genetic disease, gene editing and cancer immunotherapy. Nature Genetics 51:1645–1651.DOI:https://doi.org/10.1038/s41588-019-0517-5

Figure 1

The real-time PCR results described above have been incorporated in Figure 1, panel C, with the corresponding legend provided below (line 788).

(C) Brain cyp19a1b transcript levels in cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 6 per genotype). Mean value for cyp19a1b+/+ males was arbitrarily set to 1.

The subsequent panels have been renumbered accordingly. The entirety of the revised Figure 1.

Figure 1—figure supplement 1

The AlphaFold 3-generated structural models described above have been incorporated in Figure 1— figure supplement 1, panel C, with the corresponding legend provided below (line 811).

(C) Predicted three-dimensional structures of wild-type (left) and mutant (right) Cyp19a1b proteins. Key structural features are annotated as follows: membrane helix (blue), aromatic region (red), and heme-binding loop (orange).

The entirety of the revised Figure 1—figure supplement 1 is presented

The information on the primers used for real-time PCR has been included in Supplementary file 2.

The functional deficiency of esr2a was already addressed in the previous revision. For clarity, we have reproduced the relevant information here.

A previous study reported that female medaka lacking esr2a fail to release eggs due to oviduct atresia (Kayo et al., 2019, Sci Rep 9:8868). Similarly, in this study, some esr2a-deficient females exhibited spawning behavior but were unable to release eggs, although the sample size was limited (Δ8 line: 2/3; Δ4 line: 1/1). In contrast, this was not observed in wild-type females (Δ8 line: 0/12; Δ4 line: 0/11). These results support the effective loss of esr2a function. To incorporate this information into the manuscript, the following text has been added to the Materials and methods (line 423): “A previous study reported that esr2a-deficient female medaka cannot release eggs due to oviduct atresia (Kayo et al., 2019). Likewise, some esr2a-deficient females generated in this study, despite the limited sample size, exhibited spawning behavior but were unable to release eggs (Δ8 line: 2/3; Δ4 line: 1/1), while such failure was not observed in wild-type females (Δ8 line: 0/12; Δ4 line: 0/11). These results support the effective loss of esr2a function.”

Most experiments are weakly powered (low sample size).

This comment is essentially the same as one raised in the first review (Reviewer #3’s comment 7 on weaknesses). We acknowledge the reviewer’s concern that the histological analyses were weakly powered due to the limited sample size. In our earlier revision, we responded as follows:

Histological analyses were conducted with a relatively small sample size, as our previous experience suggested that interindividual variability in the results would not be substantial. Since significant differences were detected in many analyses, further increasing the sample size was deemed unnecessary.

The variability of the mRNA content for a same target gene between experiments (genotype comparison vs E2 treatment comparison) raises questions about the reproducibility of the data (apparent disappearance of genotype effect).

This comment is the same as one raised in the first review (Reviewer #3’s comment 8 on weaknesses), which we already addressed in our initial revision. For the reviewer’s convenience, we provide the response below:

As the reviewer pointed out, the overall area of ara expression is larger in Figure 2J than in Figure 2F. However, the relative area ratios of ara expression among brain nuclei are consistent between the two figures, indicating the reproducibility of the results. Thus, this difference is unlikely to affect the conclusions of this study.

Additionally, the differences in ara expression in pPPp and arb expression in aPPp between wild-type and cyp19a1b-deficient males appear less pronounced in Figures 2J and 2K than in Figures 2F and 2H. This is likely attributable to the smaller sample size used in the experiments for Figures 2J and 2K, resulting in less distinct differences. However, as the same genotype-dependent trends are observed in both sets of figures, the conclusion that ara and arb expression is reduced in cyp19a1b-deficient male brains remains valid.

Conclusions:

Overall, the claims regarding role of estrogens originating in the brain on male sexual behavior is supported by converging evidence from multiple mutant lines. The role of brain-derived estrogens on gene expression in the brain is weaker as are the results in females.

We appreciate the reviewer’s positive evaluation of our findings on male behavior. The concern regarding the role of brain-derived estrogens in gene expression has been addressed in our rebuttal, and the female data have been removed so that the analysis now focuses on males. The specific revisions for removing the female data are described in Response to reviewer #1’s comment 6 on weaknesses.

Recommendations For The Authors:

Reviewer #1 (Recommendations For The Authors):

The manuscript is improved slightly. I am thankful the authors addressed some concerns, but for several concerns the referees raised, the authors acknowledged them yet did not make corresponding changes to the manuscript or disagreed that they were issues at all without explanation. All reviewers had issues with the imbalanced focus on males versus females and the male aggression assay. Yet, they did not perform additional experiments or even make changes to the framing and scope of the manuscript. If the authors had removed the female data, they may have had a more cohesive story, but then they would still be left with inadequate behavior assays in the males. If the authors don't have the time or resources to perform the additional work, then they should have said so. However, the work would be incomplete relative to the claims. That is a key point here. If they change their scope and claims, the authors avoid overstating their findings. I want to see this work published because I believe it moves the field forward. But the authors need to be realistic in their interpretations of their data.

In response to this and related comments, we have removed the female data and focused the manuscript on analyses in males. The specific revisions are described in Response to reviewer #1’s comment 6 on weaknesses. Additionally, we have validated that the cyp19a1b mutation in our medaka leads to loss of function (see Response to reviewer #3’s comment 1 on weaknesses), which further strengthens the reliability of our conclusions regarding male behavior.

I agree with the reviewer who said we need to see validation of the absence of functional cyp19a1 b in the brain. However, the results from staining for the protein and performing in situ could be quizzical. Indeed, there aren't antibodies that could distinguish between aromatase a and b, and it is not uncommon for expression of a mutated gene to be normal. One approach they could do is measure aromatase activity, but they are *sort of* doing that by measuring brain E2. It's not perfect, but we teleost folks are limited in these areas. At the very least, they should show the predicted protein structure of the mutated aromatase alleles. It could show clearly that the tertiary structure is utterly absent, giving more support to the fact that their aromatase gene is non-functional.

As noted above, we have further validated the loss of cyp19a1b function by measuring cyp19a1b transcript levels in the brain and predicting the three-dimensional structure of the mutant protein. These analyses confirmed that cyp19a1b function is indeed lost, thereby increasing the reliability of our conclusions. For further details, please refer to Response to reviewer #3’s comment 1 on weaknesses.

With all of this said, the work is important, and it is possible that with a reframing of the impact of their work in the context of their findings, I could consider the work complete. I think with a proper reframing, the work is still impactful.

In accordance with this feedback, and as described above, we have reframed the manuscript by removing the female data and focusing exclusively on males. This revision clarifies the scope of our study and reinforces the support for our conclusions. For further details, please refer to Response to reviewer #1’s comment 6 on weaknesses.

(1) Clearly state in the Figure 1 legend that each data point for male aggressive behaviors represents the total # of behaviors calculated over the 4 males in each experimental tank.

In response to this comment, we have revised the legend of Figure 1K (line 797). The original legend, “(K) Total number of each aggressive act observed among cyp19a1b+/+, cyp19a1b+/−, or cyp19a1−/− males in the tank (n = 6, 7, and 5, respectively),” has been updated to “(K) Total number of each aggressive act performed by cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males. Each data point represents the sum of acts recorded for the 4 males of the same genotype in a single tank (n = 6, 7, and 5 tanks, respectively).” This clarifies that each data point reflects the total behaviors of the 4 males within each tank.

(2) The authors wrote under "Response to reviewer #1's major comment "...the development of male behaviors may require moderate neuroestrogen levels that are sufficient to induce the expression of ara and arb, but not esr2b, in the underlying neural circuitry": "This may account for the lack of aggression recovery in E2-treated cyp19a1b-deficient males in this study.".

What is meant by the latter statement? What accounts for the lack of aggression? The lack of increase in esr2b? Please clarify.

Line 365: In response to this comment, “This may account for the lack of aggression recovery in E2treated cyp19a1b-deficient males in this study.” has been revised to “Considering this, the lack of aggression recovery in E2-treated cyp19a1b-deficient males in this study may be explained by the possibility that the E2 dose used was sufficient to induce not only ara and arb but also esr2b expression in aggression-relevant circuits, which potentially suppressed aggression.”

This revision clarifies that, while moderate brain estrogen levels are sufficient to promote male behaviors via induction of ara and arb, the E2 dose used in this study may have additionally induced esr2b in circuits relevant to aggression, potentially underlying the lack of aggression recovery.

(3) This is a continuation of my comment/concern directly above. If the induction of ara and arb aren't enough, then how can, as the authors state, androgen signaling be the primary driver of these behaviors?

In response to this follow-up comment, we would like to clarify that, as described above, the lack of aggression recovery in E2-treated cyp19a1b-deficient males is not due to insufficient induction of ara and arb, but instead is likely because esr2b was also induced in aggression-relevant circuits, which may have suppressed aggression. Therefore, the concern that androgen signaling cannot be the primary driver of these behaviors is not applicable.

(4) The authors' point about sticking with the terminology for the ar genes as "ara" and "arb" is not convincing. The whole point of needing a change to match the field of neuroendocrinology as a whole (that is, across all vertebrates) is researchers, especially those with high standing like the Okubo group, adopt the new terminology. Indeed, the Okubo group is THE leader in medaka neuroendocrinology. It would go a long way if they began adopting the new terminology of "ar1" and "ar2". I understand this may be laborious to a degree, and each group can choose to use their terminology, but I'd be remiss if I didn't express my opinion that changing the terminology could help our field as a whole.

We sincerely appreciate the reviewer’s thoughtful comments regarding nomenclature consistency in vertebrate neuroendocrinology. We understand the motivation behind the suggestion to adopt ar1 and ar2. However, we consider the established nomenclature of ara and arb to be more appropriate for the following reasons.

First, adopting the ar1/ar2 nomenclature would introduce a discrepancy between gene and protein symbols. According to the NCBI International Protein Nomenclature Guidelines (Section 2B.Abbreviations and symbols;

https://www.ncbi.nlm.nih.gov/genbank/internatprot_nomenguide/), the ZFIN Zebrafish Nomenclature Conventions (Section 2. PROTEINS:https://zfin.atlassian.net/wiki/spaces/general/pages/1818394635/ZFIN+Zebrafish+Nomenclature+Con ventions), and the author guidelines of many journal

(e.g.,https://academic.oup.com/molehr/pages/Gene_And_Protein_Nomenclature), gene and protein symbols should be identical (with proteins designated in non-italic font and with the first letter capitalized). Maintaining consistency between gene and protein symbols helps avoid unnecessary confusion. The ara/arb nomenclature allows this, whereas ar1/ar2 does not.

Second, the two androgen receptor genes in teleosts are paralogs derived from the third round of wholegenome duplication that occurred early in teleost evolution. For such duplicated genes, the ZFIN Zebrafish Nomenclature Conventions (Section 1.2. Duplicated genes) recommend appending the suffixes “a” and “b” to the approved symbol of the human or mouse ortholog. This convention clearly indicates that these genes are whole-genome duplication paralogs and provides an intuitive way to represent orthologous and paralogous relationships between teleost genes and those of other vertebrates. As a result, it has been widely adopted, and we consider it logical and beneficial to apply the same principle to androgen receptors.

In light of these considerations, we respectfully maintain that the ara/arb nomenclature is more suitable for the present manuscript than the alternative ar1/ar2 system.

(5) In the discussion please discuss these potentially unexpected findings.

(a) gal was unaffected in female cyp19a1 mutants, but they exhibit mating behaviors towards females. Given gal is higher in males and these females act like females, what does this mean about the function of gal/its utility in being a male-specific marker (is it one??)?

(b) esr2b expression is higher in female cyp19a1 mutants. this is unexpected as well given esr2b is required for female-typical mating and is higher in females compared to males and E2 increases esr2b expression. please explain...well, what this means for our idea of what esr2b expression tell us.

We thank the reviewer for the insightful comments. As the female data have been removed from the manuscript, discussion of these findings in female cyp19a1b mutants is no longer necessary.

Reviewer #3 (Recommendations For The Authors):

The authors have addressed a number of answers to the reviewer's comments, notably they provided missing methodological information and rephrased the text. However, the authors have not addressed the main issues raised by the reviewers. Notably, it is regrettable that the reduced amount of brain aromatase cannot be confirmed, this seems to be the primary step when validating a new mutant. Even if protein products of the two genes may not be discriminated (which I can understand), it should be possible to evaluate the expression of a common messenger and/or peptide and confirm that aromatase expression is reduced in the brain. Since Cyp19a1b is relatively more abundant in the brain Cyp19a1a, this would strengthen the conclusion and provide confidence that the mutant indeed does silence aromatase expression in the brain. Although these short comings are acknowledged in the rebuttal letter, this is not mentioned in the discussion. Doing so would make the manuscript more transparent and clearer.

As noted in Response to reviewer #3’s comment 1 on weaknesses, we have validated the loss of Cyp19a1b function by measuring its transcript levels in the brain and predicting the three-dimensional structure of the mutant protein. These analyses confirmed that Cyp19a1b function is indeed lost, thereby increasing the reliability of our conclusions.

FigS1 - panels C&D please indicate in which tissue were hormones measured. Blood?

We thank the reviewer for pointing this out. In our study, “peripheral” refers to the caudal half of the body excluding the head and visceral organs, not blood. Accordingly, we have revised the figure legend and the description in the Materials and Methods section as follows:

Legend for Figure 1B (line 787) now reads: “Levels of E2, testosterone, and 11KT in the brain (A) and peripheral tissues (caudal half of the body) (B) of adult cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 3 per genotype).”

Materials and methods (line 431): The sentence “Total lipids were extracted from the brain and peripheral tissues (from the caudal half) of” has been revised to “Total lipids were extracted from the brain and from peripheral tissues, specifically the caudal half of the body excluding the head and visceral organs, of.”

Additional Alterations:

We have reformatted the text and supporting materials to comply with the journal’s Author Guidelines. The following changes have been made:

(1) Figures and supplementary files are now provided separately from the main text.

(2) The title page has been reformatted without any changes to its content.

(3) In-text citations have been changed from numerical references to the author–year format.

(4) Figure labels have been revised from “Fig. 1,” “Fig. S1,” etc., to “Figure 1,” “Figure 1—figure supplement 1,” etc.

(5) Table labels have been revised from “Table S1,” etc., to “Supplementary file 1,” etc.

(6) Line 324: The typo “is” has been corrected to “are”.

(7) Line 382: The section heading “Materials and Methods” has been changed to “Materials and methods” (lowercase “m”).

(8) Line 383: The Key Resources Table has been placed at the beginning of the Materials and methods section.

(9) Line 389: The sentence “Sexually mature adults (2–6 months) were used for experiments, and tissues were consistently sampled 1–5 hours after lights on.” has been revised to “Sexually mature adults (2–6 months) were used for experiments and assigned randomly to experimental groups. Tissues were consistently sampled 1–5 hours after lights on.”

(10) Line 393: The sentence “All fish were handled in accordance with the guidelines of the Institutional Animal Care and Use Committee of the University of Tokyo.” has been removed.

(11) Line 589: The following sentence has been added: “No power analysis was conducted due to the lack of relevant data; sample size was estimated based on previous studies reporting inter-individual variation in behavior and neural gene expression in medaka.”

(12) Line 598: The reference list has been reordered from numerical sequence to alphabetical order by author.

(13) In the figure legends, notations such as “A and B” have been revised to “A, B.”

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: In rodents, estrogens aromatized from androgens in the brain are essential for the development of male-typical behaviors. In many other vertebrates including humans and teleost fish, however, androgens facilitate these behaviors directly via the androgen receptor without aromatization into estrogens. Here we report that mutagenesis-derived male medaka fish lacking Cyp19a1b (a subtype of aromatase predominantly expressed in the brain) exhibit severely impaired male-typical mating and aggression, despite elevated brain androgen levels. These phenotypes can be rescued by estrogen administration, indicating that brain-derived estrogens are pivotal for male-typical behaviors even in teleosts. Our results further suggest that these estrogens facilitate male-typical behaviors by potentiating androgen action in the brain via the direct stimulation of androgen receptor transcription, thereby revealing a previously unappreciated mode of action of brain-derived estrogens. We additionally show that female fish lacking Cyp19a1b are less receptive to male courtship and conversely court other females, highlighting the significance of brain-derived estrogens in establishing sex-typical behaviors in both sexes.

Revised version: In rodents, estrogens aromatized from androgens in the brain are essential for the development of male-typical behaviors. In many other vertebrates including humans and teleost fish, however, androgens facilitate these behaviors directly via the androgen receptor without aromatization into estrogens. Here we report that mutagenesis-derived male medaka fish lacking Cyp19a1b (a subtype of aromatase predominantly expressed in the brain) exhibit severely impaired male-typical mating and aggression, despite elevated brain androgen levels. These phenotypes can be rescued by estrogen administration, indicating that brain-derived estrogens are pivotal for male-typical behaviors even in teleosts. Our results further suggest that these estrogens facilitate male-typical behaviors by potentiating androgen action in the brain via the direct stimulation of androgen receptor transcription. Taken together, these findings reveal a previously unappreciated mode of action of brain-derived estrogens in shaping male-typical behaviors.

Introduction

Earlier version: Male and female animals exhibit differences in many innate behaviors, such as mating and aggression (1). In vertebrates, these differences are driven primarily by the influence of sex steroid hormones, including estrogens and androgens. Extensive research in rodents has established that estrogens, traditionally considered “female” hormones, are critical for the development of male-typical behaviors (2–4). Specifically, androgens secreted by the testis, both perinatally and in adulthood, are converted to estrogens in the brain by the enzyme aromatase. These estrogens then act through ESR1, a subtype of the estrogen receptor (ESR), to elicit male-typical behaviors (2–4). This process, originally referred to as the “aromatization hypothesis”, is now widely acknowledged; however, the mechanisms through which brain-derived estrogens affect male-typical behavior, including the identity of their downstream targets, remain largely elusive (4, 5).

Revised version: Male and female animals exhibit differences in many innate behaviors, such as mating and aggression (Zilkha et al., 2021). In vertebrates, these differences are driven primarily by the influence of sex steroid hormones, including estrogens and androgens. Extensive research in rodents has established that estrogens, traditionally considered “female” hormones, are critical for the development of male-typical behaviors (Ogawa et al., 2020; Tsukahara et al., 2020; McCarthy, 2023). Specifically, androgens secreted by the testis, both perinatally and in adulthood, are converted to estrogens in the brain by the enzyme aromatase. These estrogens then act through ESR1, a subtype of the estrogen receptor (ESR), to elicit male-typical behaviors (Ogawa et al., 2020; Tsukahara et al., 2020; McCarthy, 2023). This process, originally referred to as the “aromatization hypothesis”, is now widely acknowledged; however, the mechanisms through which brain-derived estrogens affect male-typical behavior, including the identity of their downstream targets, remain largely elusive (McCarthy et al., 2017; McCarthy, 2023).

Earlier version: More importantly, the aromatization hypothesis seems to apply to only a limited number of species besides rodents (e.g., some birds such as zebra finches) (6). In other species such as humans, other primates, and teleost fish, testicular androgens facilitate male-typical behaviors directly through the androgen receptor (AR) without aromatization (7–9). In primates, the hypothalamic aromatization of androgens to estrogens plays a central role in female gametogenesis (10) but is not essential for male behaviors (7, 8). Notably, in teleosts, 11-ketotestosterone (11KT), which cannot be aromatized to estrogens, is the primary testicular androgen, and exogenous 11KT effectively induces male-typical courtship and aggression even in females (9, 11, 12). Therefore, it is generally assumed that, in teleosts, androgens are largely responsible for male-typical behaviors, while estrogens are dispensable for these behaviors. This is consistent with recent observations in a few teleost species that genetic ablation of AR severely impairs male-typical behaviors (13–16) and with findings in medaka fish (Oryzias latipes) that estrogens act through the ESR subtype Esr2b to prevent females from engaging in male-typical courtship (12).

Revised version: More importantly, the aromatization hypothesis seems to apply to only a limited number of species besides rodents (e.g., some birds such as zebra finches) (Balthazart, 2019). In other species such as humans, other primates, and teleost fish, testicular androgens facilitate male-typical behaviors directly through the androgen receptor (AR) without aromatization (Thornton et al., 2009; Bakker, 2022; Okubo et al., 2022). In primates, the hypothalamic aromatization of androgens to estrogens plays a central role in female gametogenesis (Terasawa, 2018) but is not essential for male behaviors (Thornton et al., 2009; Bakker, 2022). Notably, in teleosts, 11-ketotestosterone (11KT), which cannot be aromatized to estrogens, is the primary testicular androgen, and exogenous 11KT effectively induces male-typical courtship and aggression even in females (Nishiike et al., 2021; Okubo et al., 2022; Kawabata-Sakata et al., 2024). Therefore, it is generally assumed that, in teleosts, androgens are largely responsible for male-typical behaviors, while estrogens are dispensable for these behaviors. This is consistent with recent observations in a few teleost species that genetic ablation of AR severely impairs male-typical behaviors (Yong et al., 2017; Alward et al., 2020; Ogino et al., 2023; Nishiike and Okubo, 2024) and with findings in medaka fish (Oryzias latipes) that estrogens act through the ESR subtype Esr2b to prevent females from engaging in male-typical courtship (Nishiike et al., 2021).

Earlier version: Interestingly, despite these findings, adult teleost brains have extremely high levels of aromatase activity (100–1000 times higher than in rodent brains), resulting in large amounts of brain estrogens even in males (17). In teleost brains, including those of medaka, aromatase is exclusively localized in radial glial cells, in contrast to its neuronal localization in rodent brains (18–20). These observations suggest that brain-derived estrogens have a vital, but as yet undetermined, role in male teleosts. It is worth mentioning that systemic administration of estrogens and an aromatase inhibitor increased and decreased male aggression, respectively, in several teleost species, potentially reflecting the behavioral effects of brain-derived estrogens (21–24). Most teleost species have two distinct genes encoding aromatase, cyp19a1a and cyp19a1b, due to a whole-genome duplication that occurred early in teleost evolution (17, 25). These genes have undergone subfunctionalization through the partitioning of tissue-specific expression patterns: cyp19a1a is expressed predominantly in the gonad, whereas cyp19a1b is expressed in the brain (17, 25). In medaka, cyp19a1b is also expressed in the gonads, but only at a level tens to hundreds of times lower than in the brain and substantially lower than that of cyp19a1a (26, 27).

Revised version: Interestingly, despite these findings, adult teleost brains have extremely high levels of aromatase activity (100–1000 times higher than in rodent brains), resulting in large amounts of brain estrogens even in males (Diotel et al., 2018). In teleost brains, including those of medaka, aromatase is exclusively localized in radial glial cells, in contrast to its neuronal localization in rodent brains (Forlano et al., 2001; Diotel et al., 2010; Takeuchi and Okubo, 2013). These observations suggest that brain-derived estrogens have a vital, but as yet undetermined, role in male teleosts. It is worth mentioning that systemic administration of estrogens and an aromatase inhibitor increased and decreased male aggression, respectively, in several teleost species, potentially reflecting the behavioral effects of brain-derived estrogens (Hallgren et al., 2006; O’Connell and Hofmann, 2012; Huffman et al., 2013; Jalabert et al., 2015). Most teleost species have two distinct genes encoding aromatase, cyp19a1a and cyp19a1b, due to a whole-genome duplication that occurred early in teleost evolution (Diotel et al., 2018; Nagahama et al., 2021). These genes have undergone subfunctionalization through the partitioning of tissue-specific expression patterns: cyp19a1a is expressed predominantly in the gonad, whereas cyp19a1b is expressed in the brain (Diotel et al., 2018; Nagahama et al., 2021). In medaka, cyp19a1b is also expressed in the gonads, but only at a level tens to hundreds of times lower than in the brain and substantially lower than that of cyp19a1a (Okubo et al., 2011; Nakamoto et al., 2018).

In the present study, we generated cyp19a1b-deficient medaka, in which estrogen synthesis in the brain is selectively impaired while that in the gonads remains intact, in order to investigate the impact of brain-derived estrogens on male behaviors. Remarkably, the fish showed severely impaired male-typical behaviors, thus revealing the marked behavioral effects of these estrogens. Our results further suggest that these effects are mediated by potentiating androgen signaling in behaviorally relevant brain regions.

Results

cyp19a1b-deficient males exhibit severely impaired male-typical mating and aggressive behaviors

Earlier version: We generated a cyp19a1b-deficient medaka line from a mutant founder carrying a nonsense mutation in exon 4 of cyp19a1b that was identified through screening of the medaka TILLING (targeting-induced local lesions in genomes) library (28) (Fig. S1, A and B). Loss of cyp19a1b function in these fish was verified by measuring brain and peripheral levels of sex steroids. As expected, brain estradiol-17β (E2) in both male and female homozygous mutants (cyp19a1b−/−) was significantly reduced to 16% and 50%, respectively, of the levels in their wild-type (cyp19a1b+/+) siblings (P = 0.0037, males; P = 0.0092, females) (Fig. 1, A and B). In males, brain E2 in heterozygotes (cyp19a1b+/−) was also reduced to 45% of the level in wild-type siblings (P = 0.0284) (Fig. 1A), indicating a dosage effect of cyp19a1b mutation. In contrast, peripheral E2 levels were unaltered in both cyp19a1b−/− males and females (Fig. S1, C and D), consistent with the expected functioning of Cyp19a1b primarily in the brain. Strikingly, brain levels of testosterone, as opposed to E2, increased 2.2-fold in cyp19a1b−/− males relative to wild-type siblings (P = 0.0006) (Fig. 1A). Similarly, brain 11KT levels in cyp19a1b−/− males and females increased 6.2- and 1.9-fold, respectively, versus wild-type siblings (P = 0.0007, males; P = 0.0316, females) (Fig. 1, A and B). These results show that cyp19a1b-deficient fish have reduced estrogen levels coupled with increased androgen levels in the brain, confirming the loss of cyp19a1b function. They also suggest that the majority of estrogens in the male brain and half of those in the female brain are synthesized locally in the brain. In addition, peripheral 11KT levels in cyp19a1b−/− males and females increased 3.7- and 1.8-fold, respectively (P = 0.0789, males; P = 0.0118, females) (Fig. S1, C and D), indicating peripheral influence in addition to central effects.

Revised version: We generated a cyp19a1b-deficient medaka line from a mutant founder carrying a nonsense mutation in exon 4 of cyp19a1b that was identified through screening of the medaka TILLING (targeting-induced local lesions in genomes) library (Taniguchi et al., 2006) (Figure 1—figure supplement 1A and B). Loss of cyp19a1b function in these fish was verified by measuring brain and peripheral levels of sex steroids in males. As expected, brain estradiol-17β (E2) in homozygous mutants (cyp19a1b−/−) was significantly reduced to 16% of the levels in wild-type (cyp19a1b+/+) siblings (P = 0.0037) (Figure 1A). Brain E2 in heterozygotes (cyp19a1b+/−) was also reduced to 45% of wild-type levels (P = 0.0284) (Figure 1A), indicating a dosage effect of the cyp19a1b mutation. In contrast, peripheral E2 levels were unaltered in cyp19a1b−/− males (Figure 1B), consistent with the expected functioning of Cyp19a1b primarily in the brain. Strikingly, brain testosterone levels, as opposed to E2, increased 2.2-fold in cyp19a1b−/− males relative to wild-type siblings (P = 0.0006) (Figure 1A). Similarly, brain 11KT levels increased 6.2-fold (P = 0.0007) (Figure 1A). These results indicate that cyp19a1b-deficient males have reduced estrogen coupled with elevated androgen levels in the brain, confirming the loss of cyp19a1b function. They also suggest that the majority of estrogens in the male brain are synthesized locally in the brain. Peripheral 11KT levels also increased 3.7-fold in cyp19a1b−/− males (P = 0.0789) (Figure 1B), indicating peripheral influence in addition to central effects.

cyp19a1b-deficient males exhibit severely impaired male-typical mating and aggressive behaviors.

Earlier version: (A and B) Levels of E2, testosterone, and 11KT in the brain of adult cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (A) and females (B) (n = 3 per genotype and sex). (C) Set-up for testing the mating behavior of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males. (D) Latency of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 18 per genotype) to initiate each mating act toward the stimulus female. (E) Set-up for testing mating behavior using an esr2b-deficient female as the stimulus. (F) Latency of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 16, 15, and 17, respectively) to initiate each mating act toward the esr2b-deficient female. (G) Number of each mating act performed. (H) Latency of cyp19a1b+/+ and cyp19a1b−/− males and cyp19a1b+/+ females (n = 15 each) to receive courtship displays from the esr2b-deficient female. (I) Set-up for testing aggressive behavior among grouped males. (J) Total number of each aggressive act observed among cyp19a1b+/+, cyp19a1b+/−, or cyp19a1b−/− males in the tank (n = 6, 7, and 5, respectively). Statistical differences were assessed by Bonferroni’s or Dunn’s post hoc test (A, B, G, and J) and Gehan-Breslow-Wilcoxon test with Bonferroni’s correction (D, F, and H). Error bars represent SEM. *P < 0.05, **P < 0.01, ***P < 0.001.

Revised version: (A, B) Levels of E2, testosterone, and 11KT in the brain (A) and periphery (B) of adult cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 3 per genotype). (C) Brain cyp19a1b transcript levels in cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 6 per genotype). Mean value for cyp19a1b+/+ males was arbitrarily set to 1. (D) Set-up for testing the mating behavior of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males. (E) Latency of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 18 per genotype) to initiate each mating act toward the stimulus female. (F) Set-up for testing mating behavior using an esr2b-deficient female as the stimulus. (G) Latency of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males (n = 16, 15, and 17, respectively) to initiate each mating act toward the esr2b-deficient female. (H) Number of each mating act performed. (I) Latency of cyp19a1b+/+ and cyp19a1b−/− males and cyp19a1b+/+ females (n = 15 each) to receive courtship displays from the esr2b-deficient female. (J) Set-up for testing aggressive behavior among grouped males. (K) Total number of each aggressive act performed by cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males. Each data point represents the sum of acts recorded for the 4 males of the same genotype in a single tank (n = 6, 7, and 5 tanks, respectively). Statistical differences were assessed by Bonferroni’s or Dunn’s post hoc test (A, B, C, H, K) and Gehan-Breslow-Wilcoxon test with Bonferroni’s correction (E, G, I). Error bars represent SEM. *P < 0.05, **P < 0.01, ***P < 0.001.

Added in the revised version: Loss of cyp19a1b function was further confirmed by measuring cyp19a1b transcript levels in the brain and by predicting the three-dimensional structure of the mutant protein. Real-time PCR revealed that transcript levels were reduced by half in cyp19a1b+/− males and were nearly undetectable in cyp19a1b−/− males, presumably as a result of nonsense-mediated mRNA decay (Lindeboom et al., 2019) (Figure 1C). The wild-type protein, modeled by AlphaFold 3, exhibited a typical cytochrome P450 fold, including the membrane helix, aromatic region, and heme-binding loop, all arranged in the expected configuration (Figure 1—figure supplement 1C). The mutant protein, in contrast, was severely truncated, retaining only the membrane helix (Figure 1—figure supplement 1C). The absence of essential domains strongly indicates that the allele encodes a non-functional Cyp19a1b protein. Together, transcript and structural analyses consistently demonstrate that the mutation generated in this study causes a complete loss of cyp19a1b function.

Earlier version: Next, we investigated the mating behavior of cyp19a1b-deficient males (Fig. 1C). The mating behavior of medaka follows a stereotypical sequence. It begins with the male approaching and closely following the female (following). The male then performs a courtship display, rapidly swimming in a circular pattern in front of the female. If the female is receptive, the male grasps her with his fins (wrapping), culminating in the simultaneous release of eggs and sperm (spawning) (16). Because 11KT, the primary driver of male-typical behaviors in teleosts, was increased in cyp19a1b-deficient males, we predicted that these fish would engage more actively in mating. Nevertheless, cyp19a1b−/− males showed significantly longer latencies to initiate wrappings (P = 0.0033 versus cyp19a1b+/+, P = 0.0195 versus cyp19a1b+/−) and to spawn (P = 0.0051 versus cyp19a1b+/+, P = 0.0195 versus cyp19a1b+/−) (Fig. 1D). These results suggest that they are less motivated to mate, though an alternative interpretation that their cognition or sexual preference may be altered cannot be dismissed.

Revised version: Next, we investigated the mating behavior of cyp19a1b-deficient males (Figure 1D). The mating behavior of medaka follows a stereotypical sequence. It begins with the male approaching and closely following the female (following). The male then performs a courtship display, rapidly swimming in a circular pattern in front of the female. If the female is receptive, the male grasps her with his fins (wrapping), culminating in the simultaneous release of eggs and sperm (spawning) (Nishiike and Okubo, 2024). Because 11KT, the primary driver of male-typical behaviors in teleosts, was increased in cyp19a1b-deficient males, we predicted that these fish would engage more actively in mating.

Added in the revised version: Nevertheless, cyp19a1b−/− males showed significantly longer latencies to initiate wrappings (P = 0.0033 versus cyp19a1b+/+, P = 0.0195 versus cyp19a1b+/−) and to spawn (P = 0.0051 versus cyp19a1b+/+, P = 0.0195 versus cyp19a1b+/−) (Figure 1E). These results suggest that they are less motivated to mate, though an alternative interpretation that their cognition or sexual preference may be altered cannot be dismissed.

Earlier version: However, no significant differences were evident in latencies to initiate followings and courtship displays (Fig. 1D); therefore, the possibility that cyp19a1b-deficient males are less sexually attractive and less preferred by females could not be ruled out. To ascertain whether cyp19a1b-deficient males are indeed less motivated to mate, we further tested their mating behavior using esr2b-deficient females, which are unreceptive to male courtship (12), as stimulus females (Fig. 1E); this test eliminates the influence of female receptivity, facilitating an exclusive evaluation of male motivation to mate with females (16). We found that cyp19a1b−/− males showed significantly longer latencies to initiate followings (P = 0.0129 versus cyp19a1b+/+, P = 0.0060 versus cyp19a1b+/−), courtship displays (P = 0.0003 versus cyp19a1b+/+, P = 0.0006 versus cyp19a1b+/−), and wrapping attempts (P = 0.0009 versus cyp19a1b+/+, P = 0.0282 versus cyp19a1b+/−) (Fig. 1F). In addition, they exhibited significantly fewer courtship displays (P < 0.0001 versus both cyp19a1b+/+ and cyp19a1b+/−) and wrapping attempts (P < 0.0001 versus both cyp19a1b+/+ and cyp19a1b+/−) (Fig. 1G). These results thus confirmed that cyp19a1b-deficient males are less motivated to mate. In our previous study, we found that esr2b-deficient females court females preferentially over males (12). Consistent with that finding, here we observed that 12 of 15 cyp19a1b+/+ females received courtship displays from the esr2b-deficient female, as compared with only 2 of 15 cyp19a1b+/+ males (P = 0.0003) (Fig. 1H). Curiously, we further observed that 8 of 15 cyp19a1b−/− males received courtship displays from the esr2b-deficient female (P = 0.0321 versus cyp19a1b+/+ males) (Fig. 1H). Perhaps cyp19a1b−/− males are misidentified as females by esr2b-deficient females because they are reluctant to court or they exhibit some female-like behavior.

Revised version: However, no significant differences were evident in latencies to initiate followings and courtship displays (Figure 1E); therefore, the possibility that cyp19a1b-deficient males are less sexually attractive and less preferred by females could not be ruled out. To ascertain whether cyp19a1b-deficient males are indeed less motivated to mate, we further tested their mating behavior using esr2b-deficient females, which are unreceptive to male courtship (Nishiike et al., 2021), as stimulus females (Figure 1F); this test eliminates the influence of female receptivity, facilitating an exclusive evaluation of male motivation to mate with females (Nishiike and Okubo, 2024). We found that cyp19a1b−/− males showed significantly longer latencies to initiate followings (P = 0.0129 versus cyp19a1b+/+, P = 0.0060 versus cyp19a1b+/−), courtship displays (P = 0.0003 versus cyp19a1b+/+, P = 0.0006 versus cyp19a1b+/−), and wrapping attempts (P = 0.0009 versus cyp19a1b+/+, P = 0.0282 versus cyp19a1b+/−) (Figure 1G). In addition, they exhibited significantly fewer courtship displays (P < 0.0001 versus both cyp19a1b+/+ and cyp19a1b+/−) and wrapping attempts (P < 0.0001 versus both cyp19a1b+/+ and cyp19a1b+/−) (Figure 1H). These results thus confirmed that cyp19a1b-deficient males are less motivated to mate. In our previous study, we found that esr2b-deficient females court females preferentially over males (Nishiike et al., 2021). Consistent with that finding, here we observed that 12 of 15 cyp19a1b+/+ females received courtship displays from the esr2b-deficient female, as compared with only 2 of 15 cyp19a1b+/+ males (P = 0.0003) (Figure 1I). Curiously, we further observed that 8 of 15 cyp19a1b−/− males received courtship displays from the esr2b-deficient female (P = 0.0321 versus cyp19a1b+/+ males) (Figure 1I). Perhaps cyp19a1b−/− males are misidentified as females by esr2b-deficient females because they are reluctant to court or they exhibit some female-like behavior.

Earlier version: Next, we examined the intrasexual aggressive behavior of cyp19a1b-deficient males (Fig. 1I). The aggressive behavior of teleosts including medaka consists of five types of behavioral acts: chases, fin displays, circles, strikes, and bites (29). cyp19a1b−/− males exhibited all of these aggressive acts less frequently than cyp19a1b+/+ and/or cyp19a1b+/− males, with significant differences observed for chases (P = 0.0123 versus cyp19a1b+/−), fin displays (P = 0.0214 versus cyp19a1b+/+), and bites (P = 0.0015 versus cyp19a1b+/+, P = 0.0002 versus cyp19a1b+/−) (Fig. 1J). These observations demonstrate that cyp19a1b-deficient males are less aggressive toward other males.

Revised version: Next, we examined the intrasexual aggressive behavior of cyp19a1b-deficient males (Figure 1J). The aggressive behavior of teleosts including medaka consists of five types of behavioral acts: chases, fin displays, circles, strikes, and bites (Yamashita et al., 2020). cyp19a1b−/− males exhibited all of these aggressive acts less frequently than cyp19a1b+/+ and/or cyp19a1b+/− males, with significant differences observed for chases (P = 0.0123 versus cyp19a1b+/−), fin displays (P = 0.0214 versus cyp19a1b+/+), and bites (P = 0.0015 versus cyp19a1b+/+, P = 0.0002 versus cyp19a1b+/−) (Figure 1K). These observations demonstrate that cyp19a1b-deficient males are less aggressive toward other males.

Earlier version: In tilapia (Oreochromis niloticus), depletion of cyp19a1b has been reported to cause male infertility due to efferent duct obstruction (30). If this is also the case in medaka, the observed behavioral defects might be secondary to infertility, possibly due to the perception of impaired sperm release. We therefore investigated the fertilization and hatching success of embryos derived from mating between cyp19a1b−/− males and wild-type females. Their fertilization and hatching rates were 88.5% (n = 183) and 93.2% (n = 162), respectively, indicating that cyp19a1b-deficient male medaka have normal fertility.

Revised version: In tilapia (Oreochromis niloticus), depletion of cyp19a1b has been reported to cause male infertility due to efferent duct obstruction (Zhang et al., 2019). If this is also the case in medaka, the observed behavioral defects might be secondary to infertility, possibly due to the perception of impaired sperm release. We therefore investigated the fertilization and hatching success of embryos derived from mating between cyp19a1b−/− males and wild-type females. Their fertilization and hatching rates were 88.5% (n = 183) and 93.2% (n = 162), respectively, indicating that cyp19a1b-deficient male medaka have normal fertility.

Brain-derived estrogens facilitate male-typical behaviors probably by stimulating brain AR expression

Earlier version: We considered that the impaired male-typical behaviors of cyp19a1b-deficient males might be reasonably attributed to either reduced E2 or increased 11KT in the brain. The latter possibility seemed unlikely because 11KT/AR signaling strongly promotes male-typical behaviors in teleosts (9, 13–16); therefore, we tested the former possibility by examining whether E2 treatment would rescue the behavioral phenotypes of cyp19a1b-deficient males (Fig. 2A). E2 treatment, while having no effect in cyp19a1b+/+ males (Fig. 2, B and C), significantly shortened the latency to (P = 0.0005) and increased the number of (P = 0.0006) courtship displays in cyp19a1b−/− males (Fig. 2, D and E). These results suggest that reduced E2 in the brain is the primary cause of the mating defects, highlighting a pivotal role of brain-derived estrogens in male mating behavior. However, caution is warranted, as an indirect peripheral effect of bath-immersed E2 on behavior cannot be ruled out, although this is unlikely given the comparable peripheral E2 levels in cyp19a1b-deficient and wild-type males. In contrast to mating, E2 treatment was not effective in restoring aggression in cyp19a1b−/− males (Fig. S2A). It is possible that the treatment protocol used may have failed to replicate the estrogenic milieu necessary to induce aggression in males.

Revised version: We considered that the impaired male-typical behaviors of cyp19a1b-deficient males might be reasonably attributed to either reduced E2 or increased 11KT in the brain. The latter possibility seemed unlikely because 11KT/AR signaling strongly promotes male-typical behaviors in teleosts (Yong et al., 2017; Alward et al., 2020; Okubo et al., 2022; Ogino et al., 2023; Nishiike and Okubo, 2024); therefore, we tested the former possibility by examining whether E2 treatment would rescue the behavioral phenotypes of cyp19a1b-deficient males (Figure 2A). E2 treatment, while having no effect in cyp19a1b+/+ males (Figure 2B and C), significantly shortened the latency to (P = 0.0005) and increased the number of (P = 0.0006) courtship displays in cyp19a1b−/− males (Figure 2D and E). These results suggest that reduced E2 in the brain is the primary cause of the mating defects, highlighting a pivotal role of brain-derived estrogens in male mating behavior. However, caution is warranted, as an indirect peripheral effect of bath-immersed E2 on behavior cannot be ruled out, although this is unlikely given the comparable peripheral E2 levels in cyp19a1b-deficient and wild-type males. In contrast to mating, E2 treatment was not effective in restoring aggression in cyp19a1b−/− males (Figure 2—figure supplement 1A). It is possible that the treatment protocol used may have failed to replicate the estrogenic milieu necessary to induce aggression in males.

Brain-derived estrogens facilitate male-typical behaviors probably by stimulating brain AR expression.

Earlier version: (A) Set-up for testing the mating behavior of E2-treated cyp19a1b+/+ and cyp19a1b−/− males. (B) Latency of cyp19a1b+/+ males (n = 12) to initiate courtship displays toward the stimulus female before (day 0) and after (day 4) E2 treatment. (C) Number of courtship displays performed by cyp19a1b+/+ males. (D) Latency of cyp19a1b−/− males (n = 12) to initiate courtship displays before (day 0) and after (day 4) E2 treatment. (E) Number of courtship displays performed by cyp19a1b−/− males. (F) Total area of ara expression signal in each brain nucleus of cyp19a1b+/+ (n = 6 except for pPPp, where n = 5) and cyp19a1b−/− (n = 7) males. The data are displayed in two graphs for visual clarity. (G) Representative images of ara expression in the PPa, pPPp, and NVT. (H) Total area of arb expression signal in each brain nucleus of cyp19a1b+/+ (n = 6) and cyp19a1b−/− (n = 7 except for NVT, where n = 6) males. The data are displayed in two graphs for visual clarity. (I) Representative images of arb expression in the PMp, aPPp, and NPT. (J) Total area of ara expression signal in the PPa, pPPp, and NVT of cyp19a1b+/+ and cyp19a1b−/− males treated with vehicle alone or E2 (n = 5 per group except for NVT of E2-treated cyp19a1b+/+ males, where n = 4; and PPa of vehicle-treated cyp19a1b+/+, E2-treated cyp19a1b+/+, and vehicle-treated cyp19a1b−/− males, where n = 3). (K) Total area of arb expression signal in the PMp, aPPp, and NPT of cyp19a1b+/+ and cyp19a1b−/− males treated with vehicle alone or E2 (n = 5 per group except for NPT of vehicle-treated cyp19a1b+/+ males, where n = 4). Scale bars represent 50 μm. For abbreviations of brain nuclei, see Table S1. Statistical differences were assessed by Gehan-Breslow-Wilcoxon test (B and D) and unpaired t test, with Welch’s correction where appropriate (C, E, F, H, J, and K). Error bars represent SEM. *P < 0.05, **P < 0.01, ***P < 0.001.

Revised version: (A) Set-up for testing the mating behavior of E2-treated cyp19a1b+/+ and cyp19a1b−/− males. (B) Latency of cyp19a1b+/+ males (n = 12) to initiate courtship displays toward the stimulus female before (day 0) and after (day 4) E2 treatment. (C) Number of courtship displays performed by cyp19a1b+/+ males. (D) Latency of cyp19a1b−/− males (n = 12) to initiate courtship displays before (day 0) and after (day 4) E2 treatment. (E) Number of courtship displays performed by cyp19a1b−/− males. (F) Total area of ara expression signal in each brain nucleus of cyp19a1b+/+ (n = 6 except for pPPp, where n = 5) and cyp19a1b−/− (n = 7) males. The data are displayed in two graphs for visual clarity. (G) Representative images of ara expression in the PPa, pPPp, and NVT. (H) Total area of arb expression signal in each brain nucleus of cyp19a1b+/+ (n = 6) and cyp19a1b−/− (n = 7 except for NVT, where n = 6) males. The data are displayed in two graphs for visual clarity. (I) Representative images of arb expression in the PMp, aPPp, and NPT. (J) Total area of ara expression signal in the PPa, pPPp, and NVT of cyp19a1b+/+ and cyp19a1b−/− males treated with vehicle alone or E2 (n = 5 per group except for NVT of E2-treated cyp19a1b+/+ males, where n = 4; and PPa of vehicle-treated cyp19a1b+/+, E2-treated cyp19a1b+/+, and vehicle-treated cyp19a1b−/− males, where n = 3). (K) Total area of arb expression signal in the PMp, aPPp, and NPT of cyp19a1b+/+ and cyp19a1b−/− males treated with vehicle alone or E2 (n = 5 per group except for NPT of vehicle-treated cyp19a1b+/+ males, where n = 4). Scale bars represent 50 μm. For abbreviations of brain nuclei, see Supplementary file 1. Statistical differences were assessed by Gehan-Breslow-Wilcoxon test (B, D) and unpaired t test, with Welch’s correction where appropriate (C, E, F, H, J, K). Error bars represent SEM. *P < 0.05, **P < 0.01, ***P < 0.001.

Earlier version: We then considered the mechanism by which brain-derived estrogens facilitate male-typical behaviors. Our previous study showed that exogenous E2 upregulates the expression of a subtype of AR, Ara, in the medaka brain (31; note that Ara was termed Arb in this reference). We therefore hypothesized that brain-derived estrogens may facilitate male-typical behaviors by stimulating Ara expression and thereby potentiating 11KT/Ara signaling in the brain. To test this hypothesis, we first examined ara expression in the brain of cyp19a1b-deficient males by in situ hybridization analysis. Expression of ara was significantly lower in cyp19a1b−/− males than in cyp19a1b+/+ males in several preoptic and hypothalamic nuclei that are activated upon mating and/or attack in males (15), including the PPa, pPPp, and NVT (P = 0.0134, 0.0372, and 0.0008, respectively) (Fig. 2, F and G; see Table S1 for abbreviations of brain nuclei). We then performed a similar analysis for the other AR subtype, Arb. Expression of arb was significantly lower in cyp19a1b−/− than in cyp19a1b+/+ males in other preoptic and hypothalamic nuclei activated upon mating and/or attack (16), including the PMp, aPPp, and NPT (P = 0.0009, 0.0413, and 0.0021, respectively) (Fig. 2, H and I).

Revised version: We then considered the mechanism by which brain-derived estrogens facilitate male-typical behaviors. Our previous study showed that exogenous E2 upregulates the expression of a subtype of AR, Ara, in the medaka brain (Hiraki et al., 2012; note that Ara was termed Arb in this reference). We therefore hypothesized that brain-derived estrogens may facilitate male-typical behaviors by stimulating Ara expression and thereby potentiating 11KT/Ara signaling in the brain. To test this hypothesis, we first examined ara expression in the brain of cyp19a1b-deficient males by in situ hybridization analysis. Expression of ara was significantly lower in cyp19a1b−/− males than in cyp19a1b+/+ males in several preoptic and hypothalamic nuclei that are activated upon mating and/or attack in males (Nishiike and Okubo, 2024), including the PPa, pPPp, and NVT (P = 0.0134, 0.0372, and 0.0008, respectively) (Figure 2F and G, Supplementary file 1 for abbreviations of brain nuclei). We then performed a similar analysis for the other AR subtype, Arb. Expression of arb was significantly lower in cyp19a1b−/− than in cyp19a1b+/+ males in other preoptic and hypothalamic nuclei activated upon mating and/or attack (Nishiike and Okubo, 2024), including the PMp, aPPp, and NPT (P = 0.0009, 0.0413, and 0.0021, respectively) (Figure 2H and I).

Earlier version: Next, to determine whether these reductions in ara and arb expression in cyp19a1b−/− males were the result of reduced brain E2, we tested whether E2 treatment could restore the expression of the two genes. In situ hybridization revealed significantly increased ara and arb expression in most brain nuclei of E2-treated cyp19a1b−/− males as compared with vehicle-treated cyp19a1b−/− males (similar results were obtained in cyp19a1b+/+ males) (Fig. 2, J and K; Fig. S2, B and C; Fig. S3), indicating that the decreased ara and arb expression in cyp19a1b−/− males is attributable to reduced E2 levels. Taken together, these results suggest that brain-derived estrogens elicit male-typical behaviors by stimulating ara and arb expression in behaviorally relevant brain regions.

Revised version: Next, to determine whether these reductions in ara and arb expression in cyp19a1b−/− males were the result of reduced brain E2, we tested whether E2 treatment could restore the expression of the two genes. In situ hybridization revealed significantly increased ara and arb expression in most brain nuclei of E2-treated cyp19a1b−/− males as compared with vehicle-treated cyp19a1b−/− males (similar results were obtained in cyp19a1b+/+ males) (Figure 2J and K, Figure 2—figure supplement 1B and C, Figure 2—figure supplement 2), indicating that the decreased ara and arb expression in cyp19a1b−/− males is attributable to reduced E2 levels. Taken together, these results suggest that brain-derived estrogens elicit male-typical behaviors by stimulating ara and arb expression in behaviorally relevant brain regions.

cyp19a1b deficiency impairs behaviorally relevant signaling pathways downstream of ARs

Earlier version: We considered that, if our hypothesis that brain-derived estrogens facilitate male-typical behaviors by potentiating androgen/AR signaling is correct, then cyp19a1b-deficient males should have impaired activation of behaviorally relevant genes that act downstream of ARs. Two neuropeptide genes, vt (encoding vasotocin) and gal (encoding galanin), have been implicated in male-typical mating and aggressive behaviors in various vertebrates, including medaka and other teleosts (11, 16, 29, 32, 33). In medaka, subsets of neurons in the pNVT and pPMp express vt and gal, respectively, in an androgen-dependent and hence male-biased manner, and these neurons have been implicated in male-typical behaviors (11, 29, 34). We therefore studied the expression of vt and gal in the brains of cyp19a1b-deficient males.

Revised version: We considered that, if our hypothesis that brain-derived estrogens facilitate male-typical behaviors by potentiating androgen/AR signaling is correct, then cyp19a1b-deficient males should have impaired activation of behaviorally relevant genes that act downstream of ARs. Two neuropeptide genes, vt (encoding vasotocin) and gal (encoding galanin), have been implicated in male-typical mating and aggressive behaviors in various vertebrates, including medaka and other teleosts (Yokoi et al., 2015; Tripp et al., 2020; Yamashita et al., 2020; Kawabata-Sakata et al., 2024; Nishiike and Okubo, 2024). In medaka, subsets of neurons in the pNVT and pPMp express vt and gal, respectively, in an androgen-dependent and hence male-biased manner, and these neurons have been implicated in male-typical behaviors (Kawabata et al., 2012; Yamashita et al., 2020; Kawabata-Sakata et al., 2024). We therefore studied the expression of vt and gal in the brains of cyp19a1b-deficient males.

Earlier version: In situ hybridization revealed that, as expected, expression of vt in the pNVT of cyp19a1b−/− males was significantly reduced to 18% as compared with cyp19a1b+/+ males (P = 0.0040), a level comparable to that observed in females (Fig. 3, A and B). In contrast, there were no significant differences between genotypes in other brain nuclei (Fig. 3A). Similarly, expression of gal in the pPMp of cyp19a1b−/− males was reduced to 43% as compared with cyp19a1b+/+ males (P = 0.0019), while no significant differences were observed in other brain nuclei (Fig. 3, C and D). These results demonstrate that cyp19a1b deficiency severely impairs the AR-dependent activation of behaviorally relevant vt and gal expression, suggesting that brain-derived estrogens play a substantial role in activating AR signaling.

Revised version: In situ hybridization revealed that, as expected, expression of vt in the pNVT of cyp19a1b−/− males was significantly reduced to 18% as compared with cyp19a1b+/+ males (P = 0.0040) (Figure 3A and B). In contrast, there were no significant differences between genotypes in other brain nuclei (Figure 3A). Similarly, expression of gal in the pPMp of cyp19a1b−/− males was reduced to 43% as compared with cyp19a1b+/+ males (P = 0.0019), while no significant differences were observed in other brain nuclei (Figure 3C and D). These results demonstrate that cyp19a1b deficiency severely impairs the AR-dependent activation of behaviorally relevant vt and gal expression, suggesting that brain-derived estrogens play a substantial role in activating AR signaling.

cyp19a1b deficiency impairs behaviorally relevant signaling pathways downstream of ARs.

Earlier version: (A) Total area of vt expression signal in the PMp/PPa/PMm/PMg, SC/aNVT, and pNVT of cyp19a1b+/+ and cyp19a1b−/− males and females (n = 4 and 5 per genotype for males and females, respectively). (B) Representative images of vt expression in the pNVT. (C) Total area of gal expression signal in the aPMp/PPa, pPMp, PPp, and NAT/NVT/NRL of cyp19a1b+/+ and cyp19a1b−/− males and females (n = 4 except for cyp19a1b+/+ males, where n = 3). (D) Representative images of gal expression in the pPMp. Scale bars represent 50 μm. For abbreviations of brain nuclei, see Table S1. Statistical differences were assessed by unpaired t test, with Welch’s correction where appropriate (A and C). Error bars represent SEM. **P < 0.01.

Revised version: (A) Total area of vt expression signal in the PMp/PPa/PMm/PMg, SC/aNVT, and pNVT of cyp19a1b+/+ and cyp19a1b−/− males (n = 4 per genotype). (B) Representative images of vt expression in the pNVT. (C) Total area of gal expression signal in the aPMp/PPa, pPMp, PPp, and NAT/NVT/NRL of cyp19a1b+/+ and cyp19a1b−/− males (n = 3 and 4, respectively). (D) Representative images of gal expression in the pPMp. Scale bars represent 50 μm. For abbreviations of brain nuclei, see Supplementary file 1. Statistical differences were assessed by unpaired t test, with Welch’s correction where appropriate (A, C). Error bars represent SEM. **P < 0.01.

Estrogens directly stimulate the transcription of ARs through ESRs

Earlier version: The above results led us to further explore how brain-derived estrogens stimulate the expression of ara and arb. In silico analysis of the medaka ara and arb loci identified two canonical bipartite estrogen-responsive element (ERE)-like sequences in introns 1 and 2 of ara (located at positions +1699 and +2050, respectively, relative to the transcription initiation site) and three in the 5′-flanking region of arb (at positions –1272, –2180, and –3327) (Fig. 4, A and B). Thus, brain estrogens might be able to directly activate the transcription of ara and arb. To evaluate this likelihood, we conducted in vitro transcriptional activity assays in which luciferase expression was driven by genomic fragments from the ara and arb loci containing the identified ERE-like sequences.

Revised version: The above results led us to further explore how brain-derived estrogens stimulate the expression of ara and arb. In silico analysis of the medaka ara and arb loci identified two canonical bipartite estrogen-responsive element (ERE)-like sequences in introns 1 and 2 of ara (located at positions +1699 and +2050, respectively, relative to the transcription initiation site) and three in the 5′-flanking region of arb (at positions –1272, –2180, and –3327) (Figure 4A and B). Thus, brain estrogens might be able to directly activate the transcription of ara and arb. To evaluate this likelihood, we conducted in vitro transcriptional activity assays in which luciferase expression was driven by genomic fragments from the ara and arb loci containing the identified ERE-like sequences.

Estrogens directly stimulate the transcription of ARs through ESRs.

Earlier version: (A and B) Schematic of ara (A) and arb (B) loci showing the location of the canonical bipartite ERE-like sequences. Bent arrows mark the transcription initiation sites. Nucleotides of the ERE-like sequences are denoted by capital letters, and those identical to the consensus ERE (AGGTCAnnnTGACCT) are gray-shaded. (C) Ability of E2 to directly activate ara transcription. Cultured cells were transfected with a luciferase reporter construct containing a genomic fragment upstream of exon 3 of ara, together with an Esr1, Esr2a, or Esr2b expression construct. The cells were stimulated with different concentrations of E2 and luciferase activity was measured. (D) Effect of mutations in the ERE-like sequences on the E2-induced activation of ara transcription. Cultured cells were transfected with a wild-type luciferase construct or a construct carrying a mutation in the ERE-like sequence at position +1699 (mut+1699) or +2050 (mut+2050), together with an Esr1, Esr2a, or Esr2b expression construct. The cells were stimulated with or without E2, and luciferase activity was measured. (E) Ability of E2 to directly activate arb transcription. The assay described in C was performed with a luciferase construct containing a genomic fragment upstream of the first methionine codon of arb. (F) Effect of mutations in the ERE-like sequences on the E2-induced activation of arb transcription. The assay described in D was performed with luciferase constructs each carrying a mutation in the ERE-like sequence at position –1272 (mut-1272), –2180 (mut-2180), or –3327 (mut-3327). Values are expressed as a fold change relative to a control without E2 stimulation (C and E) or a control using the wild-type construct without E2 stimulation (D and F). Statistical differences were assessed by Dunnett’s post hoc test (C and E) and unpaired t test with Bonferroni–Dunn correction (D and F). Error bars represent SEM. **P < 0.01; ***P < 0.001.

Revised version: (A, B) Schematic of ara (A) and arb (B) loci showing the location of the canonical bipartite ERE-like sequences. Bent arrows mark the transcription initiation sites. Nucleotides of the ERE-like sequences are denoted by capital letters, and those identical to the consensus ERE (AGGTCAnnnTGACCT) are gray-shaded. (C) Ability of E2 to directly activate ara transcription. Cultured cells were transfected with a luciferase reporter construct containing a genomic fragment upstream of exon 3 of ara, together with an Esr1, Esr2a, or Esr2b expression construct. The cells were stimulated with different concentrations of E2 and luciferase activity was measured. (D) Effect of mutations in the ERE-like sequences on the E2-induced activation of ara transcription. Cultured cells were transfected with a wild-type luciferase construct or a construct carrying a mutation in the ERE-like sequence at position +1699 (mut+1699) or +2050 (mut+2050), together with an Esr1, Esr2a, or Esr2b expression construct. The cells were stimulated with or without E2, and luciferase activity was measured. (E) Ability of E2 to directly activate arb transcription. The assay described in C was performed with a luciferase construct containing a genomic fragment upstream of the first methionine codon of arb. (F) Effect of mutations in the ERE-like sequences on the E2-induced activation of arb transcription. The assay described in D was performed with luciferase constructs each carrying a mutation in the ERE-like sequence at position – 1272 (mut-1272), –2180 (mut-2180), or –3327 (mut-3327). Values are expressed as a fold change relative to a control without E2 stimulation (C, E) or a control using the wild-type construct without E2 stimulation (D, F). Statistical differences were assessed by Dunnett’s post hoc test (C, E) and unpaired t test with Bonferroni–Dunn correction (D, F). Error bars represent SEM. **P < 0.01; ***P < 0.001.

Earlier version: An assay using the ara genomic fragment revealed that E2 dose-dependently increased luciferase activity in the presence of any ESR subtype (Esr1, Esr2a, or Esr2b) (Fig. 4C), suggesting that E2 has the ability to directly activate ara transcription through ESRs. The introduction of a point mutation into the ERE-like sequence at position +2050 abolished the E2-induced luciferase activity in the presence of any ESR subtype, while mutation at +1699 had no such effect (Fig. 4D). These results suggest that the ERE at position +2050 is responsible for E2-induced activation of ara transcription. Because there is evidence that a single ERE half-site is sufficient to confer E2 responsiveness on several genes (35), we performed the assay using ara genomic fragments in which a mutation was introduced in either half-site of the ERE at +2050. E2 induction of luciferase activity was abrogated in both cases (Fig. S4A), suggesting that both ERE half-sites are required to confer estrogen responsiveness on ara.

Revised version: An assay using the ara genomic fragment revealed that E2 dose-dependently increased luciferase activity in the presence of any ESR subtype (Esr1, Esr2a, or Esr2b) (Figure 4C), suggesting that E2 has the ability to directly activate ara transcription through ESRs. The introduction of a point mutation into the ERE-like sequence at position +2050 abolished the E2-induced luciferase activity in the presence of any ESR subtype, while mutation at +1699 had no such effect (Figure 4D). These results suggest that the ERE at position +2050 is responsible for E2-induced activation of ara transcription. Because there is evidence that a single ERE half-site is sufficient to confer E2 responsiveness on several genes (Klinge, 2001), we performed the assay using ara genomic fragments in which a mutation was introduced in either half-site of the ERE at +2050. E2 induction of luciferase activity was abrogated in both cases (Figure 4—figure supplement 1A), suggesting that both ERE half-sites are required to confer estrogen responsiveness on ara.

Earlier version: In the assay using the arb genomic fragment, E2 also dose-dependently increased luciferase activity in the presence of any ESR subtype (Fig. 4E). The E2-induced increase was completely abolished by point mutation of the ERE-like sequence at position –3327 in the presence of any ESR subtype (Fig. 4F). A similar effect was observed for mutation of the ERE-like sequence at position –1272, but only in the presence of Esr1 and not in the presence of Esr2a or Esr2b (Fig. 4F). These results indicate that E2 can directly stimulate arb transcription, primarily through the ERE at –3327. Mutations in either half-site of this ERE eliminated induction by E2 in the presence of Esr2a or Esr2b, but not Esr1 (Fig. S4B), suggesting that both ERE half-sites are required to confer estrogen responsiveness on arb. Collectively, these experiments suggest that the transcription of ara and arb can be directly stimulated by estrogens (including brain-derived estrogens) via the binding of ESRs to canonical bipartite EREs.

Revised version: In the assay using the arb genomic fragment, E2 also dose-dependently increased luciferase activity in the presence of any ESR subtype (Figure 4E). The E2-induced increase was completely abolished by point mutation of the ERE-like sequence at position –3327 in the presence of any ESR subtype (Figure 4F). A similar effect was observed for mutation of the ERE-like sequence at position –1272, but only in the presence of Esr1 and not in the presence of Esr2a or Esr2b (Figure 4F). These results indicate that E2 can directly stimulate arb transcription, primarily through the ERE at –3327. Mutations in either half-site of this ERE eliminated induction by E2 in the presence of Esr2a or Esr2b, but not Esr1 (Figure 4—figure supplement 1B), suggesting that both ERE half-sites are required to confer estrogen responsiveness on arb. Collectively, these experiments suggest that the transcription of ara and arb can be directly stimulated by estrogens (including brain-derived estrogens) via the binding of ESRs to canonical bipartite EREs.

Brain-derived estrogens stimulate ara and arb expression in behaviorally relevant brain regions primarily through Esr2a and Esr1, respectively

Earlier version: Next, we investigated whether brain-derived estrogens can induce ara and arb expression through ESRs in vivo and, if so, which ESR subtype (Esr1, Esr2a, or Esr2b) mediates this induction. To this end, we examined ara and arb expression in the brains of males deficient for each ESR subtype by in situ hybridization. We used previously described esr1- and esr2b-deficient medaka (12, 36) and generated esr2a-deficient medaka using the CRISPR (clustered regularly interspaced short palindromic repeats) / Cas9 (CRISPR-associated protein 9) system. Two independent esr2a-deficient lines (Δ8 and Δ4) were established (Fig. S5) and used for subsequent behavioral experiments to ensure the reproducibility of the observed phenotypes.

Revised version: Next, we investigated whether brain-derived estrogens can induce ara and arb expression through ESRs in vivo and, if so, which ESR subtype (Esr1, Esr2a, or Esr2b) mediates this induction. To this end, we examined ara and arb expression in the brains of males deficient for each ESR subtype by in situ hybridization. We used previously described esr1- and esr2b-deficient medaka (Nishiike et al., 2021; Fleming et al., 2023) and generated esr2a-deficient medaka using the CRISPR (clustered regularly interspaced short palindromic repeats) / Cas9 (CRISPR-associated protein 9) system. Two independent esr2a-deficient lines (Δ8 and Δ4) were established (Figure 5—figure supplement 1) and used for subsequent behavioral experiments to ensure the reproducibility of the observed phenotypes.

Earlier version: In esr1+/+ and esr1−/− males, ara expression was not significantly different between genotypes in any brain nucleus (Fig. S6A), but arb expression was significantly lower in esr1−/− than in esr1+/+ males in the PMp and aPPp (P = 0.0111 and 0.0376, respectively) (Fig. 5, A and B). This, together with our observation that arb expression in these brain nuclei was also significantly lower in cyp19a1b−/− males (Fig. 2H), suggests that brain-derived estrogens stimulate arb expression in these nuclei primarily through Esr1. In support of this notion, double-label in situ hybridization in the wild-type male brain detected neurons coexpressing arb and esr1 in both the PMp and aPPp (Fig. S6B). In esr2a+/+ and esr2a−/− male brains (from the Δ8 line), ara expression was significantly lower in esr2a−/− than in esr2a+/+ males in the PPa, pPPp, and NVT (P = 0.0359, 0.0430, and 0.0178, respectively) (Fig. 5, C and D), while no significant difference was observed in arb expression (Fig. S6C). Considering that ara expression in these nuclei was also lower in cyp19a1b−/− males (Fig. 2F), brain-derived estrogens presumably stimulate ara expression in these nuclei mainly through Esr2a. This idea was further supported by double-label in situ hybridization, which detected neurons coexpressing ara and esr2a in the PPa, pPPp, and NVT of wild-type males (Fig. S6D). Examination of esr2b+/+ and esr2b−/− male brains showed significantly lower expression of ara in the NAT and higher expression of arb in the PPa of esr2b−/− males (Fig. S6, E to H); however, no such changes in expression were observed in cyp19a1b−/− males (Fig. 2, F and H). It therefore seems unlikely that Esr2b mediates the effect of brain estrogens on AR expression. Collectively, these results suggest that brain-derived estrogens stimulate ara and arb expression in behaviorally relevant brain regions primarily through Esr2a and Esr1, respectively.

Revised version: In esr1+/+ and esr1−/− males, ara expression was not significantly different between genotypes in any brain nucleus (Figure 5—figure supplement 2A), but arb expression was significantly lower in esr1−/− than in esr1+/+ males in the PMp and aPPp (P = 0.0111 and 0.0376, respectively) (Figure 5A and B). This, together with our observation that arb expression in these brain nuclei was also significantly lower in cyp19a1b−/− males (Figure 2H), suggests that brain-derived estrogens stimulate arb expression in these nuclei primarily through Esr1. In support of this notion, double-label in situ hybridization in the wild-type male brain detected neurons coexpressing arb and esr1 in both the PMp and aPPp (Figure 5—figure supplement 2B). In esr2a+/+ and esr2a−/− male brains (from the Δ8 line), ara expression was significantly lower in esr2a−/− than in esr2a+/+ males in the PPa, pPPp, and NVT (P = 0.0359, 0.0430, and 0.0178, respectively) (Figure 5C and D), while no significant difference was observed in arb expression (Figure 5—figure supplement 2C). Considering that ara expression in these nuclei was also lower in cyp19a1b−/− males (Figure 2F), brain-derived estrogens presumably stimulate ara expression in these nuclei mainly through Esr2a. This idea was further supported by double-label in situ hybridization, which detected neurons coexpressing ara and esr2a in the PPa, pPPp, and NVT of wild-type males (Figure 5—figure supplement 2D). Examination of esr2b+/+ and esr2b−/− male brains showed significantly lower expression of ara in the NAT and higher expression of arb in the PPa of esr2b−/− males (Figure 5—figure supplement 2E–H); however, no such changes in expression were observed in cyp19a1b−/− males (Figure 2F and H). It therefore seems unlikely that Esr2b mediates the effect of brain estrogens on AR expression. Collectively, these results suggest that brain-derived estrogens stimulate ara and arb expression in behaviorally relevant brain regions primarily through Esr2a and Esr1, respectively.

Brain-derived estrogens stimulate ara and arb expression in behaviorally relevant brain regions primarily through Esr2a and Esr1, respectively.

Earlier version: (A) Total area of arb expression signal in each brain nucleus of esr1+/+ and esr1−/− males (n = 5 per genotype). The data are displayed in two graphs for visual clarity. (B) Representative images of arb expression in the PMp and aPPp. (C) Total area of ara expression signal in each brain nucleus of esr2a+/+ and esr2a−/− males (Δ8 line; n = 5 per genotype except for NPT of esr2a+/+ males and Vs/Vp and PMp of esr2a−/− males, where n = 4). The data are displayed in two graphs for visual clarity. (D) Representative images of ara expression in the PPa, pPPp, and NVT. (E) Set-up for testing the mating behavior of esr1+/+ and esr1−/− males using an esr2b-deficient female as the stimulus. (F) Latency of esr1+/+ and esr1−/− males (n = 24 and 31, respectively) to initiate each mating act toward the stimulus female. (G) Number of each mating act performed. (H) Set-up for testing the mating behavior of esr2a+/+ and esr2a−/− males using an esr2b-deficient female as the stimulus. (I) Latency of esr2a+/+ and esr2a−/− males (Δ8 line; n = 23 and 22, respectively) to initiate each mating act toward the stimulus female. (J) Number of each mating act performed. (K and L) Total number of each aggressive act observed among esr1+/+ or esr1−/− males (n = 6 per genotype) (K) and among esr2a+/+ or esr2a−/− males (Δ8 line; n = 8 and 7, respectively) (L) in the tank. Scale bars represent 50 μm. For abbreviations of brain nuclei, see Table S1. Statistical differences were assessed by unpaired t test, with Welch’s correction where appropriate (A, C, G, J, K, and L) and Gehan-Breslow-Wilcoxon test (F and I). Error bars represent SEM. *P < 0.05, **P < 0.01.

Revised version: (A) Total area of arb expression signal in each brain nucleus of esr1+/+ and esr1−/− males (n = 5 per genotype). The data are displayed in two graphs for visual clarity. (B) Representative images of arb expression in the PMp and aPPp. (C) Total area of ara expression signal in each brain nucleus of esr2a+/+ and esr2a−/− males (Δ8 line; n = 5 per genotype except for NPT of esr2a+/+ males and Vs/Vp and PMp of esr2a−/− males, where n = 4). The data are displayed in two graphs for visual clarity. (D) Representative images of ara expression in the PPa, pPPp, and NVT. (E) Set-up for testing the mating behavior of esr1+/+ and esr1−/− males using an esr2b-deficient female as the stimulus. (F) Latency of esr1+/+ and esr1−/− males (n = 24 and 31, respectively) to initiate each mating act toward the stimulus female. (G) Number of each mating act performed. (H) Set-up for testing the mating behavior of esr2a+/+ and esr2a−/− males using an esr2b-deficient female as the stimulus. (I) Latency of esr2a+/+ and esr2a−/− males (Δ8 line; n = 23 and 22, respectively) to initiate each mating act toward the stimulus female. (J) Number of each mating act performed. (K, L) Total number of each aggressive act observed among esr1+/+ or esr1−/− males (n = 6 per genotype) (K) and among esr2a+/+ or esr2a−/− males (Δ8 line; n = 8 and 7, respectively) (L) in the tank. Scale bars represent 50 μm. For abbreviations of brain nuclei, see Supplementary file 1. Statistical differences were assessed by unpaired t test, with Welch’s correction where appropriate (A, C, G, J, K, L) and Gehan-Breslow-Wilcoxon test (F, I). Error bars represent SEM. *P < 0.05, **P < 0.01.

esr1- and esr2a-deficient males are, respectively, less motivated to mate and less aggressive

Earlier version: If Esr1 and Esr2a truly mediate the behavioral effects of brain-derived estrogens through the activation of AR expression, then esr1- and esr2a-deficient males should exhibit impaired male-typical behaviors. More specifically, given that Esr1 and Esr2a function to activate the expression of arb and ara, which are responsible for male mating and aggression, respectively (16), esr1- and esr2a-deficient males should be less motivated to mate and less aggressive, respectively.

Revised version: If Esr1 and Esr2a truly mediate the behavioral effects of brain-derived estrogens through the activation of AR expression, then esr1- and esr2a-deficient males should exhibit impaired male-typical behaviors. More specifically, given that Esr1 and Esr2a function to activate the expression of arb and ara, which are responsible for male mating and aggression, respectively (Nishiike and Okubo, 2024), esr1- and esr2a-deficient males should be less motivated to mate and less aggressive, respectively.

Earlier version: To test these ideas, we first evaluated the mating behavior of esr1-deficient males using a wild-type female as the stimulus. As expected, esr1−/− males showed a significantly longer latency to initiate followings than esr1+/+ males (P = 0.0055) (Fig. S7), suggesting that they are less motivated to mate. This was further confirmed in tests using an esr2b-deficient female as the stimulus, where esr1−/− males showed a longer latency to followings and fewer courtship displays than esr1+/+ males (P = 0.0426 and 0.0039, respectively) (Fig. 5, E to G). In contrast, no deficits were observed in the mating behavior of esr2a−/− males toward wild-type females (Fig. S8, A and B). Although esr2a−/− males of the Δ8 line showed a shorter latency to followings than their wild-type siblings in tests using a stimulus esr2b-deficient female (P = 0.0146) (Fig. 5, H to J), this was not reproduced in the Δ4 line (Fig. S8, C and D). We subsequently assessed aggressive behavior and found no defects in esr1−/− males (Fig. 5K). Conversely, esr2a−/− males from both the Δ8 and Δ4 lines exhibited significantly fewer fin displays than their wild-type siblings (P = 0.0461 and 0.0293, respectively). Circles followed a similar pattern, with a significant reduction in the Δ8 line (P = 0.0446) and a comparable but non-significant decrease in the Δ4 line (P = 0.1512) (Fig. 5L; Fig. S8E), showing less aggression. Taken together with the previous finding that esr2b-deficient males show no deficits in either mating or aggressive behavior (12), these results suggest that brain-derived estrogens can promote mating with females by stimulating arb expression through Esr1 and can increase aggression toward other males by stimulating ara expression through Esr2a. Nonetheless, behavioral deficits in esr1- and esr2a-deficient males were relatively mild as compared with cyp19a1b-deficient males (Fig. 1) and ara- and arb-deficient males (16), suggesting that a compensatory mechanism may exist between ESR subtypes.

Revised version: To test these ideas, we first evaluated the mating behavior of esr1-deficient males using a wild-type female as the stimulus. As expected, esr1−/− males showed a significantly longer latency to initiate followings than esr1+/+ males (P = 0.0055) (Figure 5—figure supplement 3), suggesting that they are less motivated to mate. This was further confirmed in tests using an esr2b-deficient female as the stimulus, where esr1−/− males showed a longer latency to followings and fewer courtship displays than esr1+/+ males (P = 0.0426 and 0.0039, respectively) (Figure 5E–G). In contrast, no deficits were observed in the mating behavior of esr2a−/− males toward wild-type females (Figure 5—figure supplement 4A and B). Although esr2a−/− males of the Δ8 line showed a shorter latency to followings than their wild-type siblings in tests using a stimulus esr2b-deficient female (P = 0.0146) (Figure 5H– J), this was not reproduced in the Δ4 line (Figure 5—figure supplement 4C and D). We subsequently assessed aggressive behavior and found no defects in esr1−/− males (Figure 5K). Conversely, esr2a−/− males from both the Δ8 and Δ4 lines exhibited significantly fewer fin displays than their wild-type siblings (P = 0.0461 and 0.0293, respectively). Circles followed a similar pattern, with a significant reduction in the Δ8 line (P = 0.0446) and a comparable but non-significant decrease in the Δ4 line (P = 0.1512) (Figure 5L, Figure 5—figure supplement 3E), showing less aggression. Taken together with the previous finding that esr2b-deficient males show no deficits in either mating or aggressive behavior (Nishiike et al., 2021), these results suggest that brain-derived estrogens can promote mating with females by stimulating arb expression through Esr1 and can increase aggression toward other males by stimulating ara expression through Esr2a. Nonetheless, behavioral deficits in esr1- and esr2a-deficient males were relatively mild as compared with cyp19a1b-deficient males (Figure 1) and ara- and arb-deficient males (Nishiike and Okubo, 2024), suggesting that a compensatory mechanism may exist between ESR subtypes.

Removed in the revised version: cyp19a1b-deficient females are less receptive to males and instead court other females

Removed in the revised version: Finally, we tested the behavioral impact of cyp19a1b deficiency in females. Tests of the mating behavior of cyp19a1b−/− females revealed that they had significantly longer latencies to receive the first wrapping from the stimulus male (P = 0.0018 versus cyp19a1b+/+, P = 0.0117 versus cyp19a1b+/−) and to spawn (P = 0.0015 versus cyp19a1b+/+) (Fig. 6, A and B). These results are consistent with a recent study in zebrafish (Danio rerio) showing that cyp19a1b-deficient females are less receptive to male courtship (37). We further found that the number of courtship displays and wrapping attempts from the male was not increased in cyp19a1b−/− females (Fig. 6C), despite their reduced receptivity. This suggests that they are not only less receptive to male courtship but may also actively avoid males, depriving them of mating opportunities. We tested this assumption by measuring the distance between fish, which confirmed that cyp19a1b−/− females spent significantly less time near males at distances of 10–20 mm (P = 0.0438) and 20–30 mm (P = 0.0002) as compared with cyp19a1b+/+ females (Fig. 6D). We therefore explored the possibility that cyp19a1b−/− females may develop a stronger sexual preference for females than for males. In tests of the mating behavior between two females, 4 of 10 cyp19a1b−/− females exhibited courtship displays toward the stimulus female, while 0 of 9 cyp19a1b+/+ females did (P = 0.0396) (Fig. 6, E and F). These results demonstrate that cyp19a1b-deficient females, having a slightly female-oriented sexual preference, are less receptive to male courtship and, conversely, court other females.

Removed in the revised version: cyp19a1b-deficient females are less receptive to males and instead court other females.

Removed in the revised version: (A) Set-up for testing the mating behavior of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− females. (B) Latency of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− females (n = 10 per genotype) to receive each mating act from the stimulus male. (C) Number of each mating act received. (D) Percentage of time cyp19a1b+/+ and cyp19a1b−/− females (n = 6 per genotype) spent at different distances from the stimulus male. (E) Set-up for testing the mating behavior of cyp19a1b+/+ and cyp19a1b−/− females toward a stimulus female. (F) Latency of cyp19a1b+/+ (n = 9) or cyp19a1b−/− (n = 10) females to initiate courtship displays toward the stimulus female. (G and H) Total area (G) and representative images (H) of npba expression signal in the Vs/Vp of cyp19a1b+/+ and cyp19a1b−/− males and females (n = 5 per genotype and sex). (I and J) Total area (I) and representative images (J) of npba expression signal in the PMm/PMg of cyp19a1b+/+ and cyp19a1b−/− males and females (n = 5 per genotype and sex). (K) Total area of esr2b expression signal in each brain nucleus of cyp19a1b+/+ or cyp19a1b−/− females (n = 5 per genotype). The data are displayed in two graphs for visual clarity. Scale bars represent 50 μm. For abbreviations of brain nuclei, see Table S1. Statistical differences were assessed by Gehan-Breslow-Wilcoxon test, with Bonferroni’s correction where appropriate (B and F), Bonferroni’s post hoc test (C), unpaired t test with Bonferroni–Dunn correction (D), and unpaired t test, with Welch’s correction where appropriate (G, I, and K). Error bars represent SEM. *P < 0.05, **P < 0.01, ***P < 0.001.

Removed in the revised version: Notably, these behavioral phenotypes of cyp19a1b-deficient females, albeit milder, closely resemble those of esr2b-deficient females (12). This led us to hypothesize that cyp19a1b deficiency prevents female-typical, and induces male-typical, behavioral phenotypes by decreasing estrogen/Esr2b signaling. If correct, cyp19a1b-deficient females should be defective in the activation of behaviorally relevant genes that act downstream of Esr2b. In medaka, subsets of neurons in the Vs/Vp and PMm/PMg express npba (encoding neuropeptide B) in an estrogen/Esr2b signaling-dependent and hence female-specific manner, and these neurons have been implicated in female sexual receptivity (12, 38, 39). We therefore examined the expression of npba in the brains of cyp19a1b-deficient females by in situ hybridization. As expected, npba expression in the Vs/Vp and PMm/PMg of cyp19a1b−/− females was significantly reduced to 33% and 38%, respectively, of that of cyp19a1b+/+ females (Vs/Vp, P = 0.0386; PMm/PMg, P = 0.0461) (Fig. 6, G to J). Expression of esr2b in these brain nuclei is also estrogen-dependent and female-specific (12, 31); therefore, we tested whether it is also attributable to brain-derived estrogens. In situ hybridization showed no significant differences in esr2b expression between cyp19a1b+/+ and cyp19a1b−/− females in the Vs/Vp and PMm/PMg, although significant differences were seen in the VM, NAT, and NVT (Fig. 6K), suggesting that esr2b expression in the brain relies on ovarian estrogens and not on brain-derived estrogens. Taken altogether, our findings indicate that brain-derived estrogens contribute substantially to the induction of female-typical, and the prevention of male-typical, behaviors in females by activating estrogen/Esr2b signaling.

Discussion

Earlier version: In this study, we found that male medaka deficient for cyp19a1b exhibit severely impaired male-typical mating and aggression. This observation was noteworthy because the fish had markedly elevated brain levels of 11KT, the primary driver of male-typical behaviors in teleosts (9, 11). Deficits in mating were rescued by estrogen administration, indicating that reduced brain estrogen levels are the primary cause of the observed mating impairment; in other words, brain-derived estrogens are pivotal at least for male-typical mating behaviors in teleosts. This was also notable because, unlike in rodents, where the necessity for brain-derived estrogens has been codified as the “aromatization hypothesis”, these estrogens have been regarded as dispensable for male behaviors in many vertebrates, including teleosts (7–9). In addition, the rescue of behavioral phenotypes by estrogen administration in adults suggests that in teleosts, unlike in rodents, brain-derived estrogens early in life is not essential for the execution of male-typical behaviors. While brain-derived estrogens are necessary for male behaviors in both rodents and teleosts, the life stages at which they exert their behavioral effects probably differ between these species. Brain aromatase activity in teleosts increases with age and, at adulthood, reaches 100– 1000 times that in rodents (9, 17). In contrast, brain aromatase activity in rodents reaches its peak during the perinatal period and thereafter declines with age (40), although it remains important for male behavior in adulthood. This variation among species may represent the activation of brain estrogen synthesis at life stages critical for sexual differentiation of behavior that are unique to each species.

Revised version: In this study, we found that male medaka deficient for cyp19a1b exhibit severely impaired male-typical mating and aggression. This observation was noteworthy because the fish had markedly elevated brain levels of 11KT, the primary driver of male-typical behaviors in teleosts (Okubo et al., 2022; Kawabata-Sakata et al., 2024). Deficits in mating were rescued by estrogen administration, indicating that reduced brain estrogen levels are the primary cause of the observed mating impairment; in other words, brain-derived estrogens are pivotal at least for male-typical mating behaviors in teleosts. This was also notable because, unlike in rodents, where the necessity for brain-derived estrogens has been codified as the “aromatization hypothesis”, these estrogens have been regarded as dispensable for male behaviors in many vertebrates, including teleosts (Thornton et al., 2009; Bakker, 2022; Okubo et al., 2022). In addition, the rescue of behavioral phenotypes by estrogen administration in adults suggests that in teleosts, unlike in rodents, brain-derived estrogens early in life are not essential for the execution of male-typical behaviors. While brain-derived estrogens are necessary for male behaviors in both rodents and teleosts, the life stages at which they exert their behavioral effects probably differ between these species. Brain aromatase activity in teleosts increases with age and, at adulthood, reaches 100– 1000 times that in rodents (Diotel et al., 2018; Okubo et al., 2022). In contrast, brain aromatase activity in rodents reaches its peak during the perinatal period and thereafter declines with age (Lephart, 1996), although it remains important for male behavior in adulthood. This variation among species may represent the activation of brain estrogen synthesis at life stages critical for sexual differentiation of behavior that are unique to each species.

Earlier version: Brain-derived estrogens serve several functions during the process of sexual differentiation of the mouse brain, including the synthesis of prostaglandin PGE2 and the activation of synaptic and neurodevelopmental genes; however, the specific mechanism whereby these estrogens affect male behaviors remains obscure (4, 5, 41). Our findings in medaka indicate that brain-derived estrogens facilitate male-typical behaviors by potentiating androgen/AR signaling in behaviorally relevant brain regions via the direct stimulation of AR transcription. More specifically, they indicate that brain-derived estrogens (1) promote mating by stimulating the transcription of arb in some preoptic and hypothalamus nuclei via Esr1, and (2) increase aggression by stimulating the transcription of ara in other preoptic and hypothalamus nuclei through Esr2a. This model of the mechanism underlying the action of brain-derived estrogens would explain the apparent contradiction that, in teleosts, androgens per se elicit male behaviors without aromatization, while estrogen synthesis in the brain are also critical for these behaviors. Our data also indicate that the two AR genes, ara and arb, are direct downstream targets of brain-derived estrogens that mediate male behaviors, only a few of which have been identified thus far. In mice, perinatal brain estrogens increase Ar expression in the bed nucleus of the stria terminalis and preoptic area, two brain regions that have been implicated in male behaviors (42). Recent evidence suggests that ESR1 binds to the regulatory genomic region of Ar in these brain regions in mice (41). Given these facts, the idea that brain-derived estrogens enhance androgen/AR signaling by directly stimulating AR transcription may apply to a wide range of species, including rodents.

Revised version: Brain-derived estrogens serve several functions during the process of sexual differentiation of the mouse brain, including the synthesis of prostaglandin PGE2 and the activation of synaptic and neurodevelopmental genes; however, the specific mechanism whereby these estrogens affect male behaviors remains obscure (McCarthy et al., 2017; Gegenhuber et al., 2022; McCarthy, 2023). Our findings in medaka indicate that brain-derived estrogens facilitate male-typical behaviors by potentiating androgen/AR signaling in behaviorally relevant brain regions via the direct stimulation of AR transcription. More specifically, they indicate that brain-derived estrogens (1) promote mating by stimulating the transcription of arb in some preoptic and hypothalamus nuclei via Esr1, and (2) increase aggression by stimulating the transcription of ara in other preoptic and hypothalamus nuclei through Esr2a. This model of the mechanism underlying the action of brain-derived estrogens would explain the apparent contradiction that, in teleosts, androgens per se elicit male behaviors without aromatization, while estrogen synthesis in the brain are also critical for these behaviors. Our data also indicate that the two AR genes, ara and arb, are direct downstream targets of brain-derived estrogens that mediate male behaviors, only a few of which have been identified thus far. In mice, perinatal brain estrogens increase Ar expression in the bed nucleus of the stria terminalis and preoptic area, two brain regions that have been implicated in male behaviors (Juntti et al., 2010). Recent evidence suggests that ESR1 binds to the regulatory genomic region of Ar in these brain regions in mice (Gegenhuber et al., 2022). Given these facts, the idea that brain-derived estrogens enhance androgen/AR signaling by directly stimulating AR transcription may apply to a wide range of species, including rodents.

Earlier version: Consistent with our results, studies in several teleost species have shown that treatment of males with an aromatase inhibitor reduces their male-typical behaviors, while estrogens exert the opposite effect (21–24). Conversely, it has been shown in various teleosts, including medaka, that treatment with exogenous estrogens attenuates male behaviors (43–47). A possible explanation for this discrepancy is that estrogens may either stimulate or suppress male-typical behaviors, depending on their concentration. All studies showing the suppressive effects of exogenous estrogens were conducted at doses higher than those used in the present study or at doses mimicking the levels typical of adult females (48). In addition, our previous study in male medaka showed that high doses of exogenous estrogens induce the expression of esr2b, which prevents male-typical mating behavior, in behaviorally relevant brain regions (12). Thus, the development of male behaviors may require moderate brain estrogen levels that are sufficient to induce the expression of ara and arb, but not esr2b, in the underlying neural circuitry. This may account for the lack of aggression recovery in E2-treated cyp19a1b-deficient males in this study. Another possibility that is not mutually exclusive is that endogenous levels of brain estrogens are sufficient to motivate males to engage in male-typical behaviors, and therefore exogenous estrogens have no further effect. This possibility is at least likely for mating behavior, as estrogen treatment facilitated mating behavior in cyp19a1b-deficient males but not in their wild-type siblings. Further studies using cyp19a1b mutants from different teleost species are needed to explore these possibilities and to determine whether the findings in medaka hold for other teleosts.

Revised version: Consistent with our results, studies in several teleost species have shown that treatment of males with an aromatase inhibitor reduces their male-typical behaviors, while estrogens exert the opposite effect (Hallgren et al., 2006; O’Connell and Hofmann, 2012; Huffman et al., 2013; Jalabert et al., 2015). Conversely, it has been shown in various teleosts, including medaka, that treatment with exogenous estrogens attenuates male behaviors (Bayley et al., 1999; Bell, 2001; Bjerselius et al., 2001; Oshima et al., 2003; Martinović et al., 2007). A possible explanation for this discrepancy is that estrogens may either stimulate or suppress male-typical behaviors, depending on their concentration. All studies showing the suppressive effects of exogenous estrogens were conducted at doses higher than those used in the present study or at doses mimicking the levels typical of adult females (Kayo et al., 2020). In addition, our previous study in male medaka showed that high doses of exogenous estrogens induce the expression of esr2b, which prevents male-typical mating behavior, in behaviorally relevant brain regions (Nishiike et al., 2021). Thus, the development of male behaviors may require moderate brain estrogen levels that are sufficient to induce the expression of ara and arb, but not esr2b, in the underlying neural circuitry. Considering this, the lack of aggression recovery in E2-treated cyp19a1b-deficient males in this study may be explained by the possibility that the E2 dose used was sufficient to induce not only ara and arb but also esr2b expression in aggression-relevant circuits, which potentially suppressed aggression. Another possibility that is not mutually exclusive is that endogenous levels of brain estrogens are sufficient to motivate males to engage in male-typical behaviors, and therefore exogenous estrogens have no further effect. This possibility is at least likely for mating behavior, as estrogen treatment facilitated mating behavior in cyp19a1b-deficient males but not in their wild-type siblings. Further studies using cyp19a1b mutants from different teleost species are needed to explore these possibilities and to determine whether the findings in medaka hold for other teleosts.

Removed in the revised version: We also found that, as compared with males, females have higher brain levels of estrogens, with half being synthesized locally and the other half supplied by the ovaries. In many brain nuclei of medaka, expression of ara and arb is comparable between the sexes or at higher levels in females (31), suggesting that brain-derived estrogens stimulate AR expression in females as well as in males. This raises the issue of why females do not display male-typical behaviors despite such high brain levels of estrogens and active estrogen synthesis in the brain. Possibly, as mentioned above, high estrogen levels in the female brain may suppress male-typical behaviors by promoting the expression of esr2b. Another possibility is that females, despite having abundant ARs in the brain, may lack sufficient amounts of androgen to activate them. This seems highly probable because the treatment of female medaka and other teleosts with 11KT as adults readily induces male-typical behaviors, including courtship and aggression (9, 11, 12).

Removed in the revised version: We found that cyp19a1b-deficient females are less receptive to male courtship, demonstrating that brain-derived estrogens influence female, as well as male, mating behavior. This finding is notable when considering that, contrary to mammals and birds, estrogens have been deemed unnecessary for female sexual receptivity in teleosts (49). This notion essentially stems from the observation in several teleost species that exogenous prostaglandin PGF2α is sufficient to restore sexual receptivity in females deprived of ovarian secretions by ovariectomy (50–52). However, we recently challenged it by showing that female medaka deficient for esr2b are completely unreceptive to males, and thus estrogens play a critical role in female receptivity (12). Given that estrogen/Esr2b signaling feminizes the neural substrates that mediate mating behavior, while PGF2α signaling triggers female sexual receptivity, the present finding provides a likely explanation for this apparent contradiction, namely, that brain-derived estrogens, rather than or in addition to ovarian-derived circulating estrogens, may function upstream of PGF2α signaling to mediate female receptivity. This scheme potentially applies to other teleost species because a reduction in female receptivity due to cyp19a1b deficiency has been also observed in zebrafish (37).

Removed in the revised version: We further found that cyp19a1b-deficient females court other females. To our knowledge, only a few neuronal genes have been identified whose depletion causes females to display male-typical behaviors in vertebrates. These include trpc2 (encoding a cation channel crucial for olfactory cue-evoked signaling in the vomeronasal organ) and dmnt3a (encoding a methyltransferase responsible for de novo DNA methylation) in mice, and esr2b in medaka (12, 53, 54). Importantly, deficiency of cyp19a1b, unlike that of esr2b, did not completely preclude female-typical mating behavior in females, but caused them to exhibit both male- and female-typical mating behaviors. This observation suggests that, in teleosts, the neural circuits mediating these behaviors coexist in the female brain and may be activated simultaneously in the absence of brain estrogens. This probably also holds true for the male brain, given that cyp19a1b-deficient males also showed impaired, but not completely abolished, male-typical behaviors and were actively courted by esr2b-deficient females. It is thus plausible that brain-derived estrogens are required to prevent the simultaneous activation of neural circuits mediating male- and female-typical behaviors, thereby contributing to the mutually exclusive presentation of these behaviors. Such effects of brain estrogens are likely to be particularly relevant in teleosts, where many species spontaneously undergo sex reversal and their sexual phenotypes, behavioral or otherwise, are highly labile throughout life (55, 56). This may account for the extremely high levels of aromatase activity in the teleost brain.

In summary, we have shown that brain-derived estrogens promote male-typical behaviors by increasing brain sensitivity to testicular androgens through the stimulation of AR expression. Our findings challenge the prevailing view that brain-derived estrogens have little effect on male-typical behaviors in species where testicular androgens elicit these behaviors directly without aromatization to estrogens in the brain and unveil a previously unappreciated mechanism of the action of brain-derived estrogens. Because these species account for the majority of vertebrates, with rodents and some birds being the only known exceptions, the mechanism of their action that we have identified in medaka may be evolutionarily ancient and widely conserved across species.

Materials and methods

Animals and cell lines

Earlier version: Wild-type d-rR strain medaka and mutant medaka deficient for cyp19a1b and esr2a (generated in this study), esr1 (36), and esr2b (12) were maintained in a recirculating system at 28°C on a 14/10-hour light/dark cycle. Three or four times a day, they were fed live brine shrimp and dry food (Otohime; Marubeni Nisshin Feed, Tokyo, Japan). Sexually mature adults (2–6 months) were used for experiments, and tissues were consistently sampled 1–5 hours after lights on. In each experiment, siblings raised under the same conditions were used as the comparison group to control for the effects of genetic diversity and environmental variation. All fish were handled in accordance with the guidelines of the Institutional Animal Care and Use Committee of the University of Tokyo.

Revised version: Wild-type d-rR strain medaka and mutant medaka deficient for cyp19a1b and esr2a (generated in this study), esr1 (Fleming et al., 2023), and esr2b (Nishiike et al., 2021) were maintained in a recirculating system at 28°C on a 14/10-hour light/dark cycle. Three or four times a day, they were fed live brine shrimp and dry food (Otohime; Marubeni Nisshin Feed, Tokyo, Japan). Sexually mature adults (2–6 months) were used for experiments and assigned randomly to experimental groups. Tissues were consistently sampled 1–5 hours after lights on. In each experiment, siblings raised under the same conditions were used as the comparison group to control for the effects of genetic diversity and environmental variation.

HEK293T and HeLa cells used in this study were confirmed to be mycoplasma-free (Biotherapy Institute of Japan, Tokyo, Japan) and authenticated by short tandem repeat profiling (National Institute of Biomedical Innovation, Osaka, Japan).

Generation of mutant medaka

Earlier version: cyp19a1b-deficient medaka were generated essentially as previously described (12). In brief, a TILLING library of 5760 chemically mutagenized medaka (28) was screened for mutations in exons 3, 4, and 5 of cyp19a1b by direct sequencing of PCR-amplified fragments. A founder (ID: 57D05) with a nonsense mutation in exon 4 (K105*) was identified (Fig. S1, A and B) and backcrossed to the wild-type d-rR strain for more than six generations to eliminate background mutations. Heterozygous males and females were intercrossed to generate wild-type, heterozygous, and homozygous siblings. All experimental fish were subjected to genotyping by РСR amplification across the mutation, followed by high-resolution melting (HRM) analysis, using the primers and probe listed in Table S2. HRM analysis was performed on the LightCycler 480 System II (Roche Diagnostics, Basel, Switzerland) using the LCGreen Plus dye (BioFire Defense, Salt Lake City, UT).

Revised version: cyp19a1b-deficient medaka were generated essentially as previously described (Nishiike et al., 2021). In brief, a TILLING library of 5760 chemically mutagenized medaka (Taniguchi et al., 2006) was screened for mutations in exons 3, 4, and 5 of cyp19a1b by direct sequencing of PCR-amplified fragments. A founder (ID: 57D05) with a nonsense mutation in exon 4 (K105*) was identified (Figure 1—figure supplement 1A and B) and backcrossed to the wild-type d-rR strain for more than six generations to eliminate background mutations. Heterozygous males and females were intercrossed to generate wild-type, heterozygous, and homozygous siblings. All experimental fish were subjected to genotyping by РСR amplification across the mutation, followed by high-resolution melting (HRM) analysis, using the primers and probe listed in Supplementary file 2. HRM analysis was performed on the LightCycler 480 System II (Roche Diagnostics, Basel, Switzerland) using the LCGreen Plus dye (BioFire Defense, Salt Lake City, UT).

Earlier version: esr2a-deficient medaka were generated by using the CRISPR/Cas9 system. A CRISPR RNA (crRNA) targeting exon 3 of esr2a (Fasmac, Kanagawa, Japan) (Fig. S5A) was injected with trans-activating crRNA (Fasmac) and Cas9 protein (Nippon Gene Co. Ltd., Tokyo, Japan) into the cytoplasm of one- or two-cell stage embryos. The resulting fish were screened for germline mutations by outcrossing to wild-type fish and testing progeny for target site mutations by T7 endonuclease I assay (57) and direct sequencing. Two founders were identified that reproducibly yielded progeny carrying frameshift mutations that eliminated the DNA- and ligand-binding domains of Esr2a: one yielded progeny carrying an 8-bp deletion (Δ8); the other yielded progeny carrying a 4-bp deletion (Δ4) (Fig. S5, A and B). The DNA- and ligand-binding domains of medaka Esr2a were identified by sequence alignment with yellow perch (Perca flavescens) Esr2a, for which these domain locations have been reported (58). The progeny were intercrossed to obtain wild-type, heterozygous, and homozygous siblings. The genotype of each experimental fish was determined by HRM analysis as described above.

Revised version: esr2a-deficient medaka were generated by using the CRISPR/Cas9 system. A CRISPR RNA (crRNA) targeting exon 3 of esr2a (Fasmac, Kanagawa, Japan) (Figure 5—figure supplement 1A) was injected with trans-activating crRNA (Fasmac) and Cas9 protein (Nippon Gene Co. Ltd., Tokyo, Japan) into the cytoplasm of one- or two-cell stage embryos. The resulting fish were screened for germline mutations by outcrossing to wild-type fish and testing progeny for target site mutations by T7 endonuclease I assay (Kim et al., 2009) and direct sequencing. Two founders were identified that reproducibly yielded progeny carrying frameshift mutations that eliminated the DNA- and ligand-binding domains of Esr2a: one yielded progeny carrying an 8-bp deletion (Δ8); the other yielded progeny carrying a 4-bp deletion (Δ4) (Figure 5—figure supplement 1A and B). The DNA- and ligand-binding domains of medaka Esr2a were identified by sequence alignment with yellow perch (Perca flavescens) Esr2a, for which these domain locations have been reported (Lynn et al., 2008). The progeny were intercrossed to obtain wild-type, heterozygous, and homozygous siblings. The genotype of each experimental fish was determined by HRM analysis as described above.

Earlier version: A previous study reported that esr2a-deficient female medaka cannot release eggs due to oviduct atresia (59). Likewise, some esr2a-deficient females generated in this study, despite the limited sample size, exhibited spawning behavior but were unable to release eggs (Δ8 line: 2/3; Δ4 line: 1/1), while such failure was not observed in wild-type females (Δ8 line: 0/12; Δ4 line: 0/11). These results support the effective loss of esr2a function.

Revised version: A previous study reported that esr2a-deficient female medaka cannot release eggs due to oviduct atresia (Kayo et al., 2019). Likewise, some esr2a-deficient females generated in this study, despite the limited sample size, exhibited spawning behavior but were unable to release eggs (Δ8 line: 2/3; Δ4 line: 1/1), while such failure was not observed in wild-type females (Δ8 line: 0/12; Δ4 line: 0/11). These results support the effective loss of esr2a function.

Measurement of sex steroid levels

Earlier version: Brain and peripheral levels of sex steroids were determined by enzyme-linked immunosorbent assay (ELISA) (12). Total lipids were extracted from the brain and peripheral tissues (from the caudal half) of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males and females with chloroform/methanol (2:1, v/v) and purified on a Sep-Pak C18 Plus Light Cartridge (Waters Corporation, Milford, MA). Tissue levels of E2, testosterone, and 11KT were measured by using Estradiol, Testosterone, and 11-Keto Testosterone ELISA kits, respectively (Cayman Chemical Company, Ann Arbor, MI).

Revised version: Brain and peripheral levels of sex steroids were determined by enzyme-linked immunosorbent assay (ELISA) (Nishiike et al., 2021). Total lipids were extracted from the brain and from peripheral tissues, specifically the caudal half of the body excluding the head and visceral organs, of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males with chloroform/methanol (2:1, v/v) and purified on a Sep-Pak C18 Plus Light Cartridge (Waters Corporation, Milford, MA). Tissue levels of E2, testosterone, and 11KT were measured by using Estradiol, Testosterone, and 11-Keto Testosterone ELISA kits, respectively (Cayman Chemical Company, Ann Arbor, MI).

Added in the revised version: Real-time PCR

Added in the revised version: Total RNA was isolated from the brains of cyp19a1b+/+, cyp19a1b+/−, and cyp19a1b−/− males using the RNeasy Plus Universal Mini Kit (Qiagen, Hilden, Germany). cDNA was synthesized with the SuperScript VILO cDNA Synthesis Kit (Thermo Fisher Scientific, Waltham, MA). Real-time PCR was performed on the LightCycler 480 System II using the LightCycler 480 SYBR Green I Master (Roche Diagnostics). Melting curve analysis was conducted to verify that a single amplicon was obtained in each sample. The β-actin gene (actb; GenBank accession number NM_001104808) was used to normalize the levels of target transcripts. The primers used for real-time PCR are shown in Supplementary file 2.

Added in the revised version: Protein structure prediction

Added in the revised version: Structural predictions of Cyp19a1b proteins were conducted using AlphaFold 3 (Abramson et al., 2024). Amino acid sequences corresponding to the wild-type allele and the mutant allele generated in this study were submitted to the AlphaFold 3 prediction server. The resulting models were visualized with PyMOL (Schrödinger, New York, NY), and key structural features, including the membrane helix, aromatic region, and heme-binding loop, were annotated.

Mating behavior test

Earlier version: Mating behavior was tested essentially as previously described (39). In brief, on the day before behavioral testing, each focal fish (cyp19a1b-, esr1-, and esr2a-deficient lines) was placed with a stimulus fish (wild-type females in the tests shown in Fig. 1, C and D, Fig. 6, E and F, Fig. S7, and Fig. S8, A and B; esr2b-deficient females in Fig. 1, E to H, Fig. 2, A to E, Fig. 5, E to J, and Fig. S8, C and D; wild-type males in Fig. 6, A to C) in a 2-liter rectangular tank with a perforated transparent partition separating them. The set-up was modified for E2-treated males, which were kept on E2 treatment and not introduced to the test tanks until the day of testing to ensure the efficacy of E2 treatment. The partition was removed 1 hour after lights on, and fish were allowed to interact for 30 min while their behavior was recorded with a digital video camera (HC-V360MS, HC-VX985M, or HC-W870M; Panasonic Corporation, Osaka, Japan). The first 15 min of the recording was used to calculate the latency to the first following, courtship display, wrapping, and spawning. The period of interaction/recording was extended to 2 hours in tests of courtship displays received from the stimulus esr2b-deficient female and in tests of mating behavior between females, because they take longer to initiate courtship (12). In tests using an esr2b-deficient female as the stimulus fish, where the latency to spawn could not be calculated because these fish were unreceptive to males and did not spawn, the sexual motivation of the focal fish was instead assessed by counting the number of courtship displays and wrapping attempts in 30 min. The number of these mating acts was also counted in tests to evaluate the receptivity of females. In tests of mating behavior between two females, the stimulus female was marked with a small notch in the caudal fin to distinguish it from the focal female.

Revised version: Mating behavior was tested essentially as previously described (Hiraki-Kajiyama et al., 2019). In brief, on the day before behavioral testing, each focal male (cyp19a1b-, esr1-, and esr2a-deficient lines) was placed with a stimulus female (wild-type females in the tests shown in Figure 1D and E, Figure 5—figure supplement 3, and Figure 5—figure supplement 4A and B; esr2b-deficient females in Figure 1F– I, Figure 2A–E, Figure 5E–J, and Figure 5—figure supplement 4C and D) in a 2-liter rectangular tank with a perforated transparent partition separating them. The set-up was modified for E2-treated males, which were kept on E2 treatment and not introduced to the test tanks until the day of testing to ensure the efficacy of E2 treatment. The partition was removed 1 hour after lights on, and fish were allowed to interact for 30 min while their behavior was recorded with a digital video camera (HC-V360MS, HC-VX985M, or HC-W870M; Panasonic Corporation, Osaka, Japan). The first 15 min of the recording was used to calculate the latency to the first following, courtship display, wrapping, and spawning. In tests using an esr2b-deficient female as the stimulus fish, the latency to spawn could not be calculated because the female was unreceptive to males and did not spawn. Therefore, the sexual motivation of the focal male was assessed by counting the number of courtship displays and wrapping attempts in 30 min. To evaluate courtship displays performed by stimulus esr2b-deficient females toward focal males, the recording period was extended to 2 hours, as these females take longer to initiate courtship (Nishiike et al., 2021). In all video analyses, the researcher was blind to the fish genotype and treatment.

Aggressive behavior test

Earlier version: To test male–male aggressive behavior (29), four males of the same genotype (cyp19a1b-, esr1-, and esr2a-deficient lines) that were unfamiliar with one another were placed together in a test tank 1 hour after lights on. After 1 min of acclimation to the tank, fish were allowed to interact for 15 min while their behavior was video-recorded as described above. The total number of each aggressive act (chase, fin display, circle, strike, and bite) displayed by the four males in the tank was counted manually from the video recordings.

Revised version: To test male–male aggressive behavior (Yamashita et al., 2020), four males of the same genotype (cyp19a1b-, esr1-, and esr2a-deficient lines) that were unfamiliar with one another were placed together in a test tank 1 hour after lights on. After 1 min of acclimation to the tank, fish were allowed to interact for 15 min while their behavior was video-recorded as described above. The total number of each aggressive act (chase, fin display, circle, strike, and bite) displayed by the four males in the tank was counted manually from the video recordings.

E2 treatment

cyp19a1b+/+ and cyp19a1b−/− males were treated with 1 ng/ml of E2 (Fujifilm Wako Pure Chemical, Osaka, Japan), which was first dissolved in 100% ethanol (vehicle), or with the vehicle alone by immersion in water for 4 days, with daily water changes to maintain the nominal concentration. The mating and aggressive behaviors of males before and after E2 treatment were evaluated as described above. The expression of ara and arb in the brain of E2- and vehicle-treated males was investigated by single-label in situ hybridization as described below. Although the exact increase in brain E2 levels following E2 treatment was not quantified, the observed positive effects on behavior and gene expression suggest that it was sufficient.

Single-label in situ hybridization

Earlier version: Digoxigenin (DIG)-labeled cRNA probes were generated by in vitro transcription using DIG RNA Labeling Mix (Roche Diagnostics), T7 RNA polymerase (Roche Diagnostics), and the following cDNA fragments: ara, nucleotide 16–1030 of GenBank NM_001122911 (1015 bp); arb, 53–1233 of NM_001104681 (1181 bp); vt, 1–845 of NM_001278891 (845 bp); gal, 5–533 of LC532140 (529 bp); npba, 1–645 of NM_001308979 (645 bp); and esr2b, 2080–3288 of XM_020713365 (1209 bp).

Revised version: Digoxigenin (DIG)-labeled cRNA probes were generated by in vitro transcription using DIG RNA Labeling Mix (Roche Diagnostics), T7 RNA polymerase (Roche Diagnostics), and the following cDNA fragments: ara, nucleotide 16–1030 of GenBank NM_001122911 (1015 bp); arb, 53–1233 of NM_001104681 (1181 bp); vt, 1–845 of NM_001278891 (845 bp); and gal, 5–533 of LC532140 (529 bp).

Earlier version: Single-label in situ hybridization was essentially done as described (60). In brief, brains dissected from males and females of the cyp19a1b-deficient line (analysis of ara, arb, vt, gal, npba, and esr2b) and males of the esr1-, esr2a-, and esr2b-deficient lines (analysis of ara and arb) were fixed in 4% paraformaldehyde and embedded in paraffin. Serial coronal sections of 10-μm thickness were hybridized with a DIG-labeled probe. Hybridization signal was detected with anti-DIG conjugated to alkaline phosphatase (RRID: AB_514497; Roche Diagnostics) and visualized using 5-bromo-4-chloro-3-indolyl phosphate/nitro blue tetrazolium substrate (Roche Diagnostics). After color development for 15 min (gal), 40 min (npba), 2 hours (vt), or overnight (ara, arb, and esr2b), sections were imaged using a VS120 virtual slide microscope (Olympus, Tokyo, Japan). The total area of signal across all relevant sections, including both hemispheres, was calculated for each brain nucleus using Olyvia software (Olympus). Images were converted to a 256-level intensity scale, and pixels with intensities from 161 to 256 were considered signals. All sections used for comparison were processed in the same batch, without corrections between samples. Medaka brain atlases (61, 62) were used to identify brain nuclei.

Revised version: Single-label in situ hybridization was essentially done as described (Kawabata-Sakata et al., 2020). In brief, male brains from the cyp19a1b-deficient line (analysis of ara, arb, vt, and gal) and from the esr1-, esr2a-, and esr2b-deficient lines (analysis of ara and arb) were fixed in 4% paraformaldehyde and embedded in paraffin. Serial coronal sections of 10-μm thickness were hybridized with a DIG-labeled probe. Hybridization signal was detected with anti-DIG conjugated to alkaline phosphatase (RRID:AB_514497; Roche Diagnostics) and visualized using 5-bromo-4-chloro-3-indolyl phosphate/nitro blue tetrazolium substrate (Roche Diagnostics). After color development for 15 min (gal), 2 hours (vt), or overnight (ara and arb), sections were imaged using a VS120 virtual slide microscope (Olympus, Tokyo, Japan). The total area of signal across all relevant sections, including both hemispheres, was calculated for each brain nucleus using Olyvia software (Olympus). Images were converted to a 256-level intensity scale, and pixels with intensities from 161 to 256 were considered signals. All sections used for comparison were processed in the same batch, without corrections between samples. Medaka brain atlases (Anken and Bourrat, 1998; Ishikawa et al., 1999) were used to identify brain nuclei.

Transcriptional activity assay

Earlier version: Each full-length coding region of medaka Esr1 (GenBank XM_020714493), Esr2a (NM_001104702), and Esr2b (XM_020713365) cDNA was amplified by PCR and inserted into the expression vector pcDNA3.1/V5-His-TOPO (Thermo Fisher Scientific, Waltham, MA). Each nucleotide sequence of the 5′-flanking region of ara and arb was retrieved from the Ensembl medaka genome assembly and analyzed for potential canonical ERE-like sequences using Jaspar (version 5.0_alpha) and Match (public version 1.0) with default settings. No ERE-like sequence was detected in ara; therefore, the gene body and 3′-flanking region were also analyzed in this case. The transcription initiation site for ara and arb was determined based on an expressed sequence tag clone deposited in National BioResource Project (NBRP) Medaka (olova36n18) and GenBank NM_001104681, respectively. Medaka bacterial artificial chromosome (BAC) clones containing the ara (clone ID: ola1-111G01) and arb (ola1-192H15) loci were obtained from NBRP Medaka. A 4530-bp genomic DNA fragment upstream of ara exon 3 (comprising 2330 bp of the 5′-flanking region, exon 1, intron 1, exon 2, intron 2, and the first 32 bp of exon 3) was amplified by PCR from the BAC clone, fused to a P2A self-cleaving peptide sequence (63) at the 3′ end, and inserted into the NheI site of the luciferase reporter vector pGL4.10 (Promega, Madison, WI). Similarly, a 3995-bp fragment upstream of the first methionine codon of arb (comprising 3815 bp of the 5′-flanking region and 180 bp of exon1) was PCR-amplified and inserted into pGL4.10. The resulting constructs were transiently transfected into HEK293T (ara) or HeLa (arb) cells, together with an internal control vector pGL4.74 (Promega) and the Esr1, Esr2a, or Esr2b expression construct using Lipofectamine LTX (Thermo Fisher Scientific). Six hours after transfection, cells were incubated for 18 hours without and with E2 at doses of 10−11 M, 10−9 M, and 10−7 M in Dulbecco’s modified Eagle’s medium (phenol red-free) containing 5% charcoal/dextran-stripped fetal bovine serum (Cytiva, Marlborough, MA). Luciferase activity was determined using the Dual-Luciferase Reporter Assay System (Promega) on the GloMax 20/20n Luminometer (Promega). All assays were conducted in duplicate or triplicate and repeated independently three times.

Revised version: Each full-length coding region of medaka Esr1 (GenBank XM_020714493), Esr2a (NM_001104702), and Esr2b (XM_020713365) cDNA was amplified by PCR and inserted into the expression vector pcDNA3.1/V5-His-TOPO (Thermo Fisher Scientific). Each nucleotide sequence of the 5′-flanking region of ara and arb was retrieved from the Ensembl medaka genome assembly and analyzed for potential canonical ERE-like sequences using Jaspar (version 5.0_alpha) and Match (public version 1.0) with default settings. No ERE-like sequence was detected in ara; therefore, the gene body and 3′-flanking region were also analyzed in this case. The transcription initiation site for ara and arb was determined based on an expressed sequence tag clone deposited in National BioResource Project (NBRP) Medaka (olova36n18) and GenBank NM_001104681, respectively. Medaka bacterial artificial chromosome (BAC) clones containing the ara (clone ID: ola1-111G01) and arb (ola1-192H15) loci were obtained from NBRP Medaka. A 4530-bp genomic DNA fragment upstream of ara exon 3 (comprising 2330 bp of the 5′-flanking region, exon 1, intron 1, exon 2, intron 2, and the first 32 bp of exon 3) was amplified by PCR from the BAC clone, fused to a P2A self-cleaving peptide sequence (Kim et al., 2011) at the 3′ end, and inserted into the NheI site of the luciferase reporter vector pGL4.10 (Promega, Madison, WI). Similarly, a 3995-bp fragment upstream of the first methionine codon of arb (comprising 3815 bp of the 5′-flanking region and 180 bp of exon1) was PCR-amplified and inserted into pGL4.10. The resulting constructs were transiently transfected into HEK293T (ara) or HeLa (arb) cells, together with an internal control vector pGL4.74 (Promega) and the Esr1, Esr2a, or Esr2b expression construct using Lipofectamine LTX (Thermo Fisher Scientific). Six hours after transfection, cells were incubated for 18 hours without and with E2 at doses of 10-11 M, 10-9 M, and 10-7 M in Dulbecco’s modified Eagle’s medium (phenol red-free) containing 5% charcoal/dextran-stripped fetal bovine serum (Cytiva, Marlborough, MA). Luciferase activity was determined using the Dual-Luciferase Reporter Assay System (Promega) on the GloMax 20/20n Luminometer (Promega). All assays were conducted in duplicate or triplicate and repeated independently three times.

Further assays were performed with luciferase reporter constructs carrying point mutations in the ERE-like sequences to identify the ERE responsible for E2 induction of ara and arb transcription. Each half-site of the responsible ERE-like sequence was mutated into a HindIII recognition site using the PrimeSTAR Mutagenesis Basal Kit (Takara Bio, Shiga, Japan). Transcriptional activity assays with these constructs were done as described above, except that a single dose of E2 (10-7 M) was used.

Double-label in situ hybridization

DIG-labeled cRNA probes for ara and arb were prepared as described above. Fluorescein-labeled cRNA probes for esr1 and esr2a were generated by in vitro transcription using Fluorescein RNA Labeling Mix (Roche Diagnostics), SP6 or T7 RNA polymerase (Roche Diagnostics), and the following cDNA fragments: esr1, nucleotides 1694–2781 of XM_020714493 (1088 bp); esr2a, 1838–2276 of NM_001104702 plus 763 bp of 3′-untranslated sequence derived from the Ensembl medaka genome assembly (439 bp).

Earlier version: Double-label in situ hybridization was done as described previously (36), with minor modifications. In brief, brains dissected from wild-type males were fixed in 4% paraformaldehyde, embedded in paraffin, and coronally sectioned at 10-μm thickness. Sections were hybridized simultaneously with the DIG-labeled ara or arb probe and fluorescein-labeled esr1 or esr2a probe. DIG was detected with a mouse anti-DIG antibody (RRID: AB_304362; Abcam, Cambridge, UK) and visualized using the Alexa Fluor 555 Tyramide SuperBoost Kit, goat anti-mouse IgG (Thermo Fisher Scientific); fluorescein was detected with an anti-fluorescein antibody conjugated to horseradish peroxidase (RRID: AB_2737388; PerkinElmer, Waltham, MA) and visualized using the TSA Plus Fluorescein System (PerkinElmer). Cell nuclei were counterstained with DAPI. Fluorescent images were acquired with a Leica TCS SP8 confocal laser scanning microscope (Leica Microsystems, Wetzlar, Germany) and the following excitation/emission wavelengths: 552/620–700 nm (Alexa Fluor 555), 488/495–545 nm (fluorescein), and 405/410–480 nm (DAPI). Cells were identified as coexpressing the two genes when Alexa Fluor 555 and fluorescein signals were clearly observed in the cytoplasm surrounding DAPI-stained nuclei, with intensities markedly stronger than the background noise.

Revised version: Double-label in situ hybridization was done as described previously (Fleming et al., 2023), with minor modifications. In brief, brains dissected from wild-type males were fixed in 4% paraformaldehyde, embedded in paraffin, and coronally sectioned at 10-μm thickness. Sections were hybridized simultaneously with the DIG-labeled ara or arb probe and fluorescein-labeled esr1 or esr2a probe. DIG was detected with a mouse anti-DIG antibody (RRID:AB_304362; Abcam, Cambridge, UK) and visualized using the Alexa Fluor 555 Tyramide SuperBoost Kit, goat anti-mouse IgG (Thermo Fisher Scientific); fluorescein was detected with an anti-fluorescein antibody conjugated to horseradish peroxidase (RRID:AB_2737388; PerkinElmer, Waltham, MA) and visualized using the TSA Plus Fluorescein System (PerkinElmer). Cell nuclei were counterstained with DAPI. Fluorescent images were acquired with a Leica TCS SP8 confocal laser scanning microscope (Leica Microsystems, Wetzlar, Germany) and the following excitation/emission wavelengths: 552/620–700 nm (Alexa Fluor 555), 488/495–545 nm (fluorescein), and 405/410–480 nm (DAPI). Cells were identified as coexpressing the two genes when Alexa Fluor 555 and fluorescein signals were clearly observed in the cytoplasm surrounding DAPI-stained nuclei, with intensities markedly stronger than the background noise.

Removed in the revised version: Measurement of spatial distances between fish

Removed in the revised version: The distance between the two fish in the test tank was measured as described elsewhere (12). On the day preceding measurement, each focal female (cyp19a1b-deficient line) and a stimulus wild-type male were placed together in a 2-liter rectangular tank with a perforated transparent partition separating them. On the following day after lights on, the fish were transferred to a 25-cm (i,d,) cylindrical test tank containing 1.8 L of dechlorinated tap water with a transparent partition separating them. One hour after lights on, the partition was removed and the behavior of the fish was video-recorded for 10 min. Distances between fish were calculated by using UMA Tracker (64).

Statistical analysis

All quantitative data were expressed as mean ± standard error of the mean (SEM). On graphs, individual data points were plotted to indicate the underlying distribution. Behavioral time-series data were analyzed using Kaplan–Meier plots with the inclusion of fish that did not exhibit the given act within the test period.

Earlier version: Statistical analyses were done using GraphPad Prism (GraphPad Software, San Diego, CA). Data between two groups were compared using unpaired two-tailed Student’s t test. Welch’s correction was applied if the F test indicated that the variance between groups was significantly different. The Bonferroni–Dunn correction was applied for multiple comparisons between two groups. To compare data among more than two groups, one-way analysis of variance (ANOVA) was performed, followed by either Bonferroni’s (comparisons among experimental groups) or Dunnett’s (comparisons between untreated and E2-treated groups in Fig. 4, C and D) post hoc test. If the Brown-Forsythe test indicated a significant difference in variance across groups, the data were instead analyzed using the non-parametric Kruskal–Wallis test, followed by Dunn’s post hoc test. Differences between Kaplan–Meier curves were tested for significance using the Gehan–Breslow–Wilcoxon test (with Bonferroni’s correction for more than two comparisons). Fish that spawned without any courtship display were excluded from the analysis of courtship display because it was not appropriate to treat them either as fish that did not perform courtship displays within the test duration or as fish that performed the first courtship display 0 seconds before spawning. Specifically, 7/18 cyp19a1b+/+, 11/18 cyp19a1b+/−, and 6/18 cyp19a1b−/− males were excluded in Fig. 1D; 6/10 cyp19a1b+/+, 3/10 cyp19a1b+/−, and 6/10 cyp19a1b−/− females were excluded in Fig. 6B; 2/23 esr1+/+ and 5/24 esr1−/− males were excluded in Fig. S7; 2/24 esr2a+/+ and 3/23 esr2a−/− males were excluded in Fig. S8A; 0/23 esr2a+/+ and 0/23 esr2a−/− males were excluded in Fig. S8B. All data points were included in the analyses and no outliers were defined.

Revised version: Statistical analyses were done using GraphPad Prism (GraphPad Software, San Diego, CA). Data between two groups were compared using unpaired two-tailed Student’s t test. Welch’s correction was applied if the F test indicated that the variance between groups was significantly different. The Bonferroni–Dunn correction was applied for multiple comparisons between two groups. To compare data among more than two groups, one-way analysis of variance (ANOVA) was performed, followed by either Bonferroni’s (comparisons among experimental groups) or Dunnett’s (comparisons between untreated and E2-treated groups in Figure 4C and D) post hoc test. If the Brown-Forsythe test indicated a significant difference in variance across groups, the data were instead analyzed using the non-parametric Kruskal–Wallis test, followed by Dunn’s post hoc test. Differences between Kaplan–Meier curves were tested for significance using the Gehan–Breslow–Wilcoxon test (with Bonferroni’s correction for more than two comparisons). Fish that spawned without any courtship display were excluded from the analysis of courtship display because it was not appropriate to treat them either as fish that did not perform courtship displays within the test duration or as fish that performed the first courtship display 0 seconds before spawning. Specifically, 7/18 cyp19a1b+/+, 11/18 cyp19a1b+/−, and 6/18 cyp19a1b−/− males were excluded in Figure 1E; 2/23 esr1+/+ and 5/24 esr1−/− males were excluded in Figure 5—figure supplement 3; 2/24 esr2a+/+ and 3/23 esr2a−/− males were excluded in Figure 5— figure supplement 4A; 0/23 esr2a+/+ and 0/23 esr2a−/− males were excluded in Figure 5—figure supplement 4B. All data points were included in the analyses and no outliers were defined.

Added in the revised version: No power analysis was conducted due to the lack of relevant data; sample size was estimated based on previous studies reporting inter-individual variation in behavior and neural gene expression in medaka.